[1682395251] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.25%) too high
[1682395371] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.71%) too high
[1682395431] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.24%) too high
[1682395791] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.26%) too high
[1682395852] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (96.37%) too high
[1682396152] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.27%) too high
[1682396213] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (93.13%) too high
[1682396633] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.50%) too high
[1682396693] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.69%) too high
[1682397053] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.11%) too high
[1682397114] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (95.96%) too high
[1682397234] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.02%) too high
[1682397295] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (94.51%) too high
[1682397895] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.48%) too high
[1682397955] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.14%) too high
[1682398075] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.80%) too high
[1682398136] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (95.87%) too high
[1682398196] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.09%) too high
[1682398256] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.94%) too high
[1682398376] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.49%) too high
[1682398436] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (96.90%) too high
[1682398586] HOST ALERT: FW Blueshell;DOWN;SOFT;1;CRITICAL - Network Unreachable (190.121.5.68)
[1682398586] HOST ALERT: Blueshell;DOWN;SOFT;1;CRITICAL - Network Unreachable (190.121.5.67)
[1682398590] SERVICE ALERT: FW Blueshell;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398600] SERVICE ALERT: Blueshell;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398650] HOST ALERT: Blueshell;DOWN;SOFT;2;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398650] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;2;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398710] HOST ALERT: Blueshell;DOWN;SOFT;3;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398710] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;3;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398770] HOST ALERT: Blueshell;DOWN;SOFT;4;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398770] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;4;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398830] HOST ALERT: Blueshell;DOWN;SOFT;5;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398830] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;5;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398890] HOST ALERT: Blueshell;DOWN;SOFT;6;CRITICAL - Time to live exceeded (190.121.5.67)
[1682398890] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;6;CRITICAL - Time to live exceeded (190.121.5.68)
[1682398953] SERVICE ALERT: FW Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 25.26 ms
[1682398954] HOST ALERT: Blueshell;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 27.21 ms
[1682398954] HOST ALERT: FW Blueshell;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 27.25 ms
[1682398956] SERVICE ALERT: Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 27.27 ms
[1682399456] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.19%) too high
[1682399517] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (95.74%) too high
[1682399577] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.05%) too high
[1682399637] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.13%) too high
[1682399757] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.08%) too high
[1682399817] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.34%) too high
[1682399937] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.50%) too high
[1682399997] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 1.07MB (0.85%), Average OUT: 80.78MB (64.62%)Total RX: 8687.53GBytes, Total TX: 22467.60GBytes
[1682401138] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1682429030] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1682429150] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1682429623] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682429624] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682429624] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682429625] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682429751] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682429756] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682429756] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682429758] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682430001] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682430002] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682430002] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682430133] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682430135] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682430135] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682430601] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682430734] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682432834] HOST ALERT: Electric Works;DOWN;SOFT;1;CRITICAL - Network Unreachable (190.171.184.81)
[1682432844] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;1;CRITICAL - Network Unreachable (190.171.184.82)
[1682432864] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;1;CRITICAL - Network Unreachable (190.171.184.81)
[1682432875] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1682432924] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1682432934] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1682432984] HOST ALERT: Electric Works;DOWN;SOFT;3;CRITICAL - Network Unreachable (190.171.184.81)
[1682433024] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1682433074] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1682433114] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1682433134] HOST ALERT: Electric Works;DOWN;SOFT;5;CRITICAL - Network Unreachable (190.171.184.81)
[1682433174] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;5;CRITICAL - Network Unreachable (190.171.184.82)
[1682433224] HOST ALERT: Electric Works;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1682433254] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;6;CRITICAL - Network Unreachable (190.171.184.82)
[1682433314] HOST ALERT: Electric Works;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1682433344] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1682433345] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 32.93 ms
[1682433349] HOST ALERT: Electric Works;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 33.00 ms
[1682433355] SERVICE ALERT: FW Electric Works;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 36.09 ms
[1682433357] HOST ALERT: FW Electric Works;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 35.85 ms
[1682433390] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.81 ms
[1682443362] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;WARNING;SOFT;1;WARNING - IN bandwidth (88.97%) too high
[1682443378] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (88.67%) too high
[1682443483] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 55.24MB (44.19%), Average OUT: 649.84KB (0.52%)Total RX: 397.74GBytes, Total TX: 398.82GBytes
[1682443498] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 45.64MB (36.51%), Average OUT: 44.30MB (35.44%)Total RX: 8718.58GBytes, Total TX: 24878.62GBytes
[1682449734] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1682449869] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;(No output on stdout) stderr:
[1682463932] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;UNKNOWN;SOFT;1;NSClient - ERROR: No performance data from command: check_drivesize
[1682464036] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682464037] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682464038] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682464055] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;SOFT;2;connect to address 10.0.0.12 and port 12489: Connection refused
[1682464062] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1682464065] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.12 port 5666: Connection refused
[1682464155] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682464157] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;SOFT;2;Memory usage: total:23551.21 MB - used: 1860.11 MB (8%) - free: 21691.11 MB (92%)
[1682464157] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;SOFT;2;CPU Load 0% (5 min average)
[1682464158] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1682464175] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;SOFT;3;c:\ - total: 119.40 Gb - used: 85.31 Gb (71%) - free 34.09 Gb (29%)
[1682464182] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 178.45 Gb (17%) - free 845.53 Gb (83%)
[1682464245] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682467494] SERVICE ALERT: SRV04AD01.cass.cl;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.10)
[1682467497] HOST ALERT: SRV04AD01.cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.10)
[1682467536] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682467561] HOST ALERT: SRV04AD01.cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.22 ms
[1682467618] SERVICE ALERT: SRV04AD01.cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.13 ms
[1682467656] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.005 sec. response time
[1682475070] SERVICE ALERT: SRVADMIN.cass.pmc;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 1.33 ms
[1682475077] SERVICE ALERT: SRVADMIN.cass.pmc;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682475080] SERVICE ALERT: SRVADMIN.cass.pmc;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682475080] SERVICE ALERT: SRVADMIN.cass.pmc;Almacenamiento;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1682475081] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.109 port 5666: No route to host
[1682475081] HOST ALERT: SRVADMIN.cass.pmc;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.109)
[1682475094] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;CRITICAL;HARD;1;connect to address 10.0.0.109 and port 12489: No route to host
[1682475144] HOST ALERT: SRVADMIN.cass.pmc;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.109)
[1682475193] SERVICE ALERT: SRVADMIN.cass.pmc;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.109)
[1682475200] SERVICE ALERT: SRVADMIN.cass.pmc;Tiempo Activo;CRITICAL;HARD;3;connect to address 10.0.0.109 and port 12489: No route to host
[1682475203] SERVICE ALERT: SRVADMIN.cass.pmc;CPU;CRITICAL;HARD;3;connect to address 10.0.0.109 and port 12489: No route to host
[1682475203] SERVICE ALERT: SRVADMIN.cass.pmc;Almacenamiento;CRITICAL;HARD;3;connect to address 10.0.0.109 and port 12489: No route to host
[1682475206] HOST ALERT: SRVADMIN.cass.pmc;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.109)
[1682475269] HOST ALERT: SRVADMIN.cass.pmc;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.109)
[1682475332] HOST ALERT: SRVADMIN.cass.pmc;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.109)
[1682475396] HOST ALERT: SRVADMIN.cass.pmc;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.83 ms
[1682475434] SERVICE ALERT: SRVADMIN.cass.pmc;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.92 ms
[1682475440] SERVICE ALERT: SRVADMIN.cass.pmc;Almacenamiento;OK;HARD;1;c:\ - total: 299.66 Gb - used: 154.74 Gb (52%) - free 144.91 Gb (48%)
[1682475440] SERVICE ALERT: SRVADMIN.cass.pmc;CPU;OK;HARD;1;CPU Load 0% (5 min average)
[1682475440] SERVICE ALERT: SRVADMIN.cass.pmc;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 6 minute(s)
[1682475451] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:16123.56 MB - used: 1685.48 MB (10%) - free: 14438.08 MB (90%)
[1682475471] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682475561] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682475651] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Actualizacion;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1682475651] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682475867] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682476888] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Actualizacion;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1682480942] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.