[1676264688] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.38, 0.13, 0.05
[1676264808] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.25, 0.18, 0.08
[1676264928] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.16, 0.17, 0.09
[1676265048] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.08, 0.14, 0.09
[1676285565] SERVICE ALERT: Clinica PV;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1676285566] SERVICE ALERT: FW Clinica PV;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1676285578] HOST ALERT: Clinica PV;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1676285590] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1676285620] HOST ALERT: Clinica PV;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1676285668] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1676285710] HOST ALERT: Clinica PV;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1676285715] SERVICE ALERT: Clinica PV;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676285716] SERVICE ALERT: FW Clinica PV;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676285758] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1676285800] HOST ALERT: Clinica PV;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1676285820] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;4;CRITICAL - Network Unreachable (200.54.66.26)
[1676285880] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;5;CRITICAL - Network Unreachable (200.54.66.26)
[1676285890] HOST ALERT: Clinica PV;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1676285940] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;6;CRITICAL - Network Unreachable (200.54.66.26)
[1676285980] HOST ALERT: Clinica PV;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1676286000] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;7;CRITICAL - Network Unreachable (200.54.66.26)
[1676286060] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;8;CRITICAL - Network Unreachable (200.54.66.26)
[1676286068] SERVICE ALERT: Clinica PV;Ping;OK;HARD;1;PING OK - Packet loss = 79%, RTA = 38.77 ms
[1676286068] HOST ALERT: Clinica PV;UP;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 39.15 ms
[1676286110] SERVICE ALERT: FW Clinica PV;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 38.04 ms
[1676286114] HOST ALERT: FW Clinica PV;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 38.80 ms
[1676286170] SERVICE ALERT: FW Clinica PV;Ping;CRITICAL;SOFT;1;CRITICAL - Network Unreachable (200.54.66.26)
[1676286170] HOST ALERT: FW Clinica PV;DOWN;SOFT;1;CRITICAL - Network Unreachable (200.54.66.26)
[1676286190] HOST ALERT: Clinica PV;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1676286190] HOST FLAPPING ALERT: FW Clinica PV;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1676286190] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;2;CRITICAL - Network Unreachable (200.54.66.26)
[1676286195] SERVICE ALERT: Clinica PV;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1676286250] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;3;CRITICAL - Network Unreachable (200.54.66.26)
[1676286280] HOST ALERT: Clinica PV;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1676286290] SERVICE ALERT: FW Clinica PV;Ping;CRITICAL;HARD;3;CRITICAL - Network Unreachable (200.54.66.26)
[1676286310] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;4;CRITICAL - Network Unreachable (200.54.66.26)
[1676286370] HOST ALERT: Clinica PV;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1676286370] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;5;CRITICAL - Network Unreachable (200.54.66.26)
[1676286430] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;6;CRITICAL - Network Unreachable (200.54.66.26)
[1676286460] HOST ALERT: Clinica PV;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1676286490] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;7;CRITICAL - Network Unreachable (200.54.66.26)
[1676286550] HOST ALERT: Clinica PV;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1676286550] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;8;CRITICAL - Network Unreachable (200.54.66.26)
[1676286611] HOST ALERT: FW Clinica PV;UNREACHABLE;SOFT;9;CRITICAL - Network Unreachable (200.54.66.26)
[1676286640] HOST ALERT: Clinica PV;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1676286701] HOST ALERT: FW Clinica PV;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1676286730] HOST ALERT: Clinica PV;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1676286820] HOST ALERT: Clinica PV;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1676286851] HOST FLAPPING ALERT: FW Clinica PV;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1676286910] HOST ALERT: Clinica PV;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1676287000] HOST ALERT: Clinica PV;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1676291084] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676291205] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 59.16KB (0.05%), Average OUT: 491.52KB (0.39%)Total RX: 259.45GBytes, Total TX: 274.84GBytes
[1676291956] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1676292078] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1676292199] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1676294874] HOST ALERT: Clinica PV;UP;HARD;1;PING OK - Packet loss = 66%, RTA = 38.15 ms
[1676294874] SERVICE ALERT: Clinica PV;Ping;OK;HARD;1;PING OK - Packet loss = 50%, RTA = 37.78 ms
[1676294875] HOST ALERT: FW Clinica PV;UP;HARD;1;PING OK - Packet loss = 58%, RTA = 38.13 ms
[1676294875] SERVICE ALERT: FW Clinica PV;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 38.81 ms
[1676295500] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1676295501] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676295620] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1676295621] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 5.06KB (0.00%), Average OUT: 7.86KB (0.01%)Total RX: 22.57GBytes, Total TX: 22.37GBytes
[1676302742] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1676302862] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 359 processes
[1676302983] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 355 processes
[1676303043] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 326 processes
[1676313235] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676313355] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 43.35KB (0.03%), Average OUT: 479.26KB (0.38%)Total RX: 261.19GBytes, Total TX: 288.82GBytes
[1676313691] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676313812] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 6.79KB (0.01%), Average OUT: 6.65KB (0.01%)Total RX: 22.95GBytes, Total TX: 22.60GBytes
[1676314460] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1676314580] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1676324858] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676324979] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1676327290] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.
[1676333235] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1676333355] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1676333475] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1676333715] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 83% (5 min average)
[1676334172] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 6.93 ms
[1676334296] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.04 ms
[1676337315] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1676337435] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1676337555] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1676337735] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 78% (5 min average)