[1674961229] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 526.18KB (0.42%), Average OUT: 102.55MB (82.04%)Total RX: 9080.27GBytes, Total TX: 19298.14GBytes
[1674961829] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (93.83%) too high
[1674961949] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.05%) too high
[1674962070] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (89.83%) too high
[1674962310] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.71%) too high
[1674962370] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1674962370] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.60%) too high
[1674962730] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.01%) too high
[1674962791] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (95.75%) too high
[1674962911] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.05%) too high
[1674962972] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (90.14%) too high
[1674964112] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1674964412] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.08%) too high
[1674964473] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.72%) too high
[1674964593] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.04%) too high
[1674964654] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1674964654] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.32%) too high
[1674965254] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.38%) too high
[1674965314] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.07%) too high
[1674965974] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.59%) too high
[1674966095] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (95.65%) too high
[1674966215] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.05%) too high
[1674966276] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (96.54%) too high
[1674966457] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 837.57KB (0.67%), Average OUT: 103.08MB (82.46%)Total RX: 9093.84GBytes, Total TX: 19901.58GBytes
[1674966516] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.31%) too high
[1674966636] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 853.35KB (0.68%), Average OUT: 102.96MB (82.37%)Total RX: 9094.00GBytes, Total TX: 19920.52GBytes
[1674966757] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.09%) too high
[1674966877] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 1.33MB (1.06%), Average OUT: 100.38MB (80.30%)Total RX: 9094.27GBytes, Total TX: 19945.89GBytes
[1674966997] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.99%) too high
[1674967118] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 653.73KB (0.52%), Average OUT: 105.40MB (84.32%)Total RX: 9094.52GBytes, Total TX: 19972.14GBytes
[1674967238] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.70%) too high
[1674967358] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (85.74%) too high
[1674967479] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (85.14%) too high
[1674967600] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 588.84KB (0.47%), Average OUT: 95.13MB (76.11%)Total RX: 9094.84GBytes, Total TX: 20023.52GBytes
[1674967659] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (94.77%) too high
[1674967779] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.40%) too high
[1674967900] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 639.17KB (0.51%), Average OUT: 103.73MB (82.98%)Total RX: 9095.05GBytes, Total TX: 20056.74GBytes
[1674968020] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.16%) too high
[1674968140] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 651.52KB (0.52%), Average OUT: 105.30MB (84.24%)Total RX: 9095.21GBytes, Total TX: 20082.45GBytes
[1674968320] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.44%) too high
[1674968440] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.05%) too high
[1674968561] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (85.22%) too high
[1674968621] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 609.21KB (0.49%), Average OUT: 102.28MB (81.82%)Total RX: 9095.52GBytes, Total TX: 20133.92GBytes
[1674968650] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.28, 0.09, 0.04
[1674968770] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.26, 0.16, 0.07
[1674968890] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.17, 0.16, 0.08
[1674968950] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;HARD;1;OK - load average per CPU: 0.12, 0.15, 0.08
[1674969041] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.14%) too high
[1674969162] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 13.50MB (10.80%), Average OUT: 101.36MB (81.09%)Total RX: 9102.27GBytes, Total TX: 20185.32GBytes
[1674969462] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.90%) too high
[1674969582] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 17.94MB (14.35%), Average OUT: 99.56MB (79.65%)Total RX: 9108.86GBytes, Total TX: 20227.28GBytes
[1674970122] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.81%) too high
[1674970242] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1674970242] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 15.99MB (12.79%), Average OUT: 104.68MB (83.74%)Total RX: 9119.17GBytes, Total TX: 20295.11GBytes
[1674970422] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (91.18%) too high
[1674970542] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.13%) too high
[1674970663] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (91.74%) too high
[1674971504] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 651.32KB (0.52%), Average OUT: 105.82MB (84.66%)Total RX: 9122.36GBytes, Total TX: 20435.14GBytes
[1674971923] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.44%) too high
[1674972034] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.51)
[1674972043] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.82%) too high
[1674972050] SERVICE ALERT: FW Gore Los Lagos;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.51)
[1674972079] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.50)
[1674972079] HOST ALERT: FW Gore Los Lagos;UNREACHABLE;SOFT;2;CRITICAL - Time to live exceeded (190.121.2.51)
[1674972081] SERVICE ALERT: Gore Los Lagos;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.50)
[1674972114] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 4.80 ms
[1674972118] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 5.69 ms
[1674972143] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.80 ms
[1674972145] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.94 ms
[1674972164] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 632.08KB (0.51%), Average OUT: 101.42MB (81.13%)Total RX: 9122.77GBytes, Total TX: 20500.90GBytes
[1674972344] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.55%) too high
[1674972465] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 657.53KB (0.53%), Average OUT: 104.61MB (83.69%)Total RX: 9122.96GBytes, Total TX: 20532.01GBytes
[1674972645] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (90.07%) too high
[1674972766] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 953.06KB (0.76%), Average OUT: 97.31MB (77.85%)Total RX: 9123.20GBytes, Total TX: 20562.86GBytes
[1674975116] SERVICE ALERT: FW Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 6.32 ms
[1674975240] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 5.75 ms
[1674978854] SERVICE ALERT: Clinica PV;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 35.92 ms
[1674978979] SERVICE ALERT: Clinica PV;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 35.93 ms
[1674988216] SERVICE ALERT: Pasmar;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 37%, RTA = 15.37 ms
[1674988340] SERVICE ALERT: Pasmar;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 25.87 ms
[1674988946] SERVICE ALERT: FW Clinica PV;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 55.72 ms
[1674989070] SERVICE ALERT: FW Clinica PV;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 47.10 ms
[1675016524] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1675016533] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ CPU Load;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1675016533] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ HHD C;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1675016533] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ Usemem;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1675016533] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1675016553] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.88 ms
[1675016560] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ HHD U;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1675016644] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1675016653] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1675016653] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ Usemem;OK;SOFT;2;Memory usage: total:23551.21 MB - used: 1843.50 MB (8%) - free: 21707.72 MB (92%)
[1675016653] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ HHD C;OK;SOFT;2;c:\ - total: 119.40 Gb - used: 79.37 Gb (66%) - free 40.03 Gb (34%)
[1675016680] SERVICE ALERT: SRV01USER_cass.cl;NsClient++ HHD U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 480.45 Gb (47%) - free 543.53 Gb (53%)
[1675019356] SERVICE ALERT: Qnap NAS;Sysinfo NAS;CRITICAL;SOFT;1;CRITICAL: SNMP to 10.0.0.28 is not available or wrong community string
[1675019476] SERVICE ALERT: Qnap NAS;Sysinfo NAS;OK;SOFT;2;NAS Cloud, Model TS-879U-RP, Firmware 4.3.6, Max HD number 8, No. Volume 1
[1675027693] SERVICE ALERT: FW ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 34.09 ms
[1675027817] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 33.58 ms
[1675028183] SERVICE ALERT: FW ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 33.69 ms
[1675028307] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 33.25 ms
[1675028913] SERVICE ALERT: FW ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 33.79 ms
[1675029037] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 34.01 ms
[1675030374] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1675030436] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1675032096] SERVICE ALERT: FW Cass Entel;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 35.53 ms
[1675032220] SERVICE ALERT: FW Cass Entel;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 36.15 ms
[1675032283] SERVICE ALERT: FW ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 35.97 ms
[1675032407] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 36.18 ms
[1675032895] SERVICE ALERT: ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 37.58 ms
[1675033019] SERVICE ALERT: ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 38.71 ms
[1675035286] SERVICE ALERT: FW Cass Entel;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 35.83 ms
[1675035411] SERVICE ALERT: FW Cass Entel;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 35.32 ms
[1675036797] SERVICE ALERT: FW Cass Entel;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 33.54 ms
[1675036921] SERVICE ALERT: FW Cass Entel;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 34.28 ms
[1675038384] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675038515] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available