[1701084171] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701084301] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 205.19KB (0.16%), Average OUT: 11.65MB (9.32%)Total RX: 530.05GBytes, Total TX: 801.28GBytes
[1701084318] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701084439] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 3.59KB (0.00%), Average OUT: 5.01KB (0.00%)Total RX: 4.80GBytes, Total TX: 3.61GBytes
[1701084851] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701084972] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 212.43KB (0.17%), Average OUT: 11.33MB (9.06%)Total RX: 530.19GBytes, Total TX: 809.04GBytes
[1701085644] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1701085764] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;OK;SOFT;2;SNMP OK - Timeticks: (38065786) 4 days, 9:44:17.86
[1701088031] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1701088049] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701088151] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1701088169] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 48.72KB (0.04%), Average OUT: 7.37KB (0.01%)Total RX: 5.01GBytes, Total TX: 3.65GBytes
[1701088787] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1701088907] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1701090983] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (207850631211.68%) too high
[1701091113] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;CRITICAL - IN bandwidth (3413.45%) too high
[1701091233] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 5.83MB (4.67%), Average OUT: 5.82MB (4.66%)Total RX: 559.51GBytes, Total TX: 851.82GBytes
[1701098979] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1701099100] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 46.85KB (0.04%), Average OUT: 39.66KB (0.03%)Total RX: 5.23GBytes, Total TX: 3.81GBytes
[1701099318] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701099439] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 60.28KB (0.05%), Average OUT: 517.36KB (0.41%)Total RX: 39.70GBytes, Total TX: 189.09GBytes
[1701099462] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1701099463] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701099582] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1701099584] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 337.29B (0.00%), Average OUT: 110.53B (0.00%)Total RX: 583.61GBytes, Total TX: 876.14GBytes
[1701106971] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701107091] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 5.23KB (0.00%), Average OUT: 8.47KB (0.01%)Total RX: 5.40GBytes, Total TX: 3.93GBytes
[1701107797] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1701107868] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1701107975] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1701107985] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 1.04 ms
[1701107988] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1701107991] HOST ALERT: FW ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1701107994] SERVICE ALERT: FW Clinica PV Movistar;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 35.75 ms
[1701107994] SERVICE ALERT: ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 37.23 ms
[1701108028] HOST ALERT: FW ADL Diagnostic;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 36.78 ms
[1701108095] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;OK;SOFT;2;SNMP OK - Timeticks: (40298814) 4 days, 15:56:28.14
[1701108109] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1701108118] SERVICE ALERT: FW Clinica PV Movistar;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 35.90 ms
[1701108118] SERVICE ALERT: ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 37.35 ms
[1701108120] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1701108602] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1701108734] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1701112675] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701112795] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 37.97KB (0.03%), Average OUT: 156.23KB (0.12%)Total RX: 585.02GBytes, Total TX: 878.88GBytes
[1701114921] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1701115041] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1701116311] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1701116431] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1701117245] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1701117365] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 329.02B (0.00%), Average OUT: 106.95B (0.00%)Total RX: 585.55GBytes, Total TX: 896.81GBytes
[1701118335] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (113518421474.43%) too high
[1701118456] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;CRITICAL - IN bandwidth (3861.92%) too high
[1701118581] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 262.28B (0.00%), Average OUT: 89.48B (0.00%)Total RX: 585.55GBytes, Total TX: 896.81GBytes
[1701133664] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 98% (5 min average)
[1701133784] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;SOFT;2;CPU Load 93% (5 min average)
[1701134084] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 97% (5 min average)
[1701134087] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.67%) too high
[1701134204] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1701134207] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 12.86MB (10.29%), Average OUT: 100.01MB (80.01%)Total RX: 2747.22GBytes, Total TX: 3148.01GBytes
[1701134324] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1701134538] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1701134610] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1701134984] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 84% (5 min average)
[1701137264] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1701137384] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1701137504] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1701137700] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1701137775] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1701138164] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 93% (5 min average)