[1674270393] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;1;Memory usage: total:10228.13 MB - used: 8824.99 MB (86%) - free: 1403.14 MB (14%)
[1674270513] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;OK;SOFT;2;Memory usage: total:10228.13 MB - used: 8572.01 MB (84%) - free: 1656.12 MB (16%)
[1674277427] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.24, 0.06, 0.02
[1674277547] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;SOFT;2;OK - load average per CPU: 0.08, 0.07, 0.03
[1674282167] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.14, 0.13, 0.10
[1674282288] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.13, 0.13, 0.10
[1674282408] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.13, 0.13, 0.11
[1674297888] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.42, 0.23, 0.17
[1674297948] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.20, 0.21, 0.17
[1674298488] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;HARD;1;OK - load average per CPU: 0.00, 0.04, 0.10
[1674345813] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;1;Memory usage: total:10228.13 MB - used: 8893.01 MB (87%) - free: 1335.12 MB (13%)
[1674345933] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;OK;SOFT;2;Memory usage: total:10811.29 MB - used: 8476.09 MB (78%) - free: 2335.20 MB (22%)
[1674349413] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;1;Memory usage: total:10235.28 MB - used: 9169.74 MB (90%) - free: 1065.54 MB (10%)
[1674349533] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;2;Memory usage: total:11266.12 MB - used: 9620.69 MB (85%) - free: 1645.43 MB (15%)
[1674349653] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;OK;SOFT;3;Memory usage: total:11266.12 MB - used: 9312.02 MB (83%) - free: 1954.10 MB (17%)
[1674353073] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;1;Memory usage: total:10242.11 MB - used: 8781.15 MB (86%) - free: 1460.96 MB (14%)
[1674353193] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;SOFT;2;Memory usage: total:10242.11 MB - used: 8824.20 MB (86%) - free: 1417.91 MB (14%)
[1674353313] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;WARNING;HARD;3;Memory usage: total:10242.11 MB - used: 8966.29 MB (88%) - free: 1275.82 MB (12%)
[1674353373] SERVICE ALERT: Veeam2020.cass.pmc;NsClient++ Usemem;OK;HARD;1;Memory usage: total:10242.11 MB - used: 8510.86 MB (83%) - free: 1731.25 MB (17%)