[1677899488] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1677899609] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1677906299] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.36, 0.12, 0.05
[1677906419] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.21, 0.15, 0.07
[1677906539] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.16, 0.15, 0.08
[1677906599] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.08, 0.13, 0.08
[1677910739] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.13, 0.14, 0.10
[1677910859] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1677910979] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1677926759] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.33, 0.19, 0.15
[1677927239] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.04, 0.20, 0.20
[1677927959] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.02, 0.04, 0.10
[1677941610] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677941611] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677941630] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677941741] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677941743] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677941761] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677955364] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677955496] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677967811] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677967943] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available