[1677387838] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.24, 0.07, 0.03
[1677387958] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.22, 0.14, 0.06
[1677388078] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;3;OK - load average per CPU: 0.15, 0.14, 0.07
[1677391219] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.51)
[1677391223] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.50)
[1677391235] SERVICE ALERT: FW Gore Los Lagos;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.51)
[1677391258] SERVICE ALERT: Gore Los Lagos;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.50)
[1677391283] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.96 ms
[1677391287] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.48 ms
[1677391299] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 5.85 ms
[1677391322] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.64 ms
[1677416971] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1677417091] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1677417211] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1677417331] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 92% (5 min average)
[1677428431] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 96% (5 min average)
[1677428551] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1677428671] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1677429811] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 81% (5 min average)