[1688443214] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1688443334] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688446814] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.13. Check system logs on 10.0.2.13
[1688446934] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688450414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688450534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688455814] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1688455934] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688459414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1688459534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688468414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688468534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688470814] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688470934] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688473783] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1688473903] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 4.48MB (3.59%), Average OUT: 7.19MB (5.75%)Total RX: 752.51GBytes, Total TX: 1469.45GBytes
[1688474080] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688474212] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688474401] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688474534] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688475103] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1688475223] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1688475584] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.74 ms
[1688475715] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.79 ms
[1688476421] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688476553] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688476801] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688476933] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688477354] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688477408] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 28%, RTA = 0.98 ms
[1688477485] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.96 ms
[1688477610] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 1.04 ms
[1688478247] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (85.18%) too high
[1688478367] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 32.53MB (26.02%), Average OUT: 31.25MB (25.00%)Total RX: 29766.40GBytes, Total TX: 67703.95GBytes
[1688479140] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 1.07 ms
[1688479273] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.97 ms
[1688479471] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1688479591] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 26.99KB (0.02%), Average OUT: 11.82KB (0.01%)Total RX: 7.41GBytes, Total TX: 8.39GBytes
[1688480010] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688480063] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 16%, RTA = 1.06 ms
[1688481414] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1688481534] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 272.76B (0.00%), Average OUT: 106.17B (0.00%)Total RX: 761.51GBytes, Total TX: 1481.18GBytes
[1688481614] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688481734] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688481883] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.86 ms
[1688481913] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688481987] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 66%, RTA = 2.82 ms
[1688482018] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.95 ms
[1688482228] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.96 ms
[1688482353] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 1.05 ms
[1688482863] SERVICE ALERT: SRVDNS;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1688482987] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.19 ms
[1688483133] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688483194] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 0.96 ms
[1688483677] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.80 ms
[1688483805] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 44%, RTA = 1.05 ms
[1688484416] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688484440] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 76%, RTA = 0.88 ms
[1688485035] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 1.03 ms
[1688485160] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.85 ms
[1688485530] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688485550] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688485598] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 1.07 ms
[1688485599] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 75%, RTA = 0.75 ms
[1688486327] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 0.97 ms
[1688486439] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688486443] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688486444] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688486456] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.91 ms
[1688486589] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688486591] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688486594] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688486721] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1688486726] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1688487263] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 1.11 ms
[1688487394] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.87 ms
[1688488121] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688488253] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688488801] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688488814] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688488933] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688488934] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688490124] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.92 ms
[1688490251] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.92 ms
[1688490574] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688490638] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.91 ms
[1688490641] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.93 ms
[1688490774] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 64%, RTA = 1.07 ms
[1688490904] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688490924] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 93%, RTA = 1.02 ms
[1688490959] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.94 ms
[1688490965] HOST FLAPPING ALERT: SRVDNS;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1688490965] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 28%, RTA = 0.94 ms
[1688491169] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688491222] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 76%, RTA = 0.95 ms
[1688491284] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.94 ms
[1688491292] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688491349] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 16%, RTA = 0.85 ms
[1688491417] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.94 ms
[1688491675] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1688491765] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688491807] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688491825] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688491867] SERVICE FLAPPING ALERT: SRVDNS;Ping;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1688491867] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 0.98 ms
[1688491915] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688491927] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688491961] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1688491968] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 37%, RTA = 0.87 ms
[1688492051] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688492104] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 82%, RTA = 0.97 ms
[1688492138] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 0.64 ms
[1688492146] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 50%, RTA = 0.94 ms
[1688492221] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 1.06 ms
[1688492251] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688492342] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688492371] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 0.81 ms
[1688492432] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688492468] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 37%, RTA = 0.85 ms
[1688492484] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 70%, RTA = 0.97 ms
[1688492558] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688492605] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 80%, RTA = 1.51 ms
[1688492624] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 28%, RTA = 0.97 ms
[1688492708] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688492722] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 77%, RTA = 0.95 ms
[1688492730] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 0.93 ms
[1688492791] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.68 ms
[1688492916] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.79 ms
[1688492992] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1688493062] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 68%, RTA = 1.06 ms
[1688493186] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.90 ms
[1688493197] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688493276] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688493336] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688493366] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688493456] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688493490] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.98 ms
[1688493495] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 16%, RTA = 0.78 ms
[1688493696] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 0.69 ms
[1688493700] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688493786] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688493845] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 83%, RTA = 0.97 ms
[1688493876] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688493966] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688493966] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688494001] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.99 ms
[1688494008] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 37%, RTA = 1.02 ms
[1688494146] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.73 ms
[1688494151] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688494236] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688494296] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.77 ms
[1688494326] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688494416] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688494476] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688494506] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688494527] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 77%, RTA = 0.89 ms
[1688494557] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688494596] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 90%, RTA = 0.85 ms
[1688494647] HOST ALERT: SRVDNS;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1688494702] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 70%, RTA = 0.94 ms
[1688494732] HOST ALERT: SRVDNS;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1688494776] HOST FLAPPING ALERT: SRVDNS;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1688494776] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688494822] HOST ALERT: SRVDNS;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1688494836] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 1.26 ms
[1688494872] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 0.72 ms
[1688494885] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 64%, RTA = 0.89 ms
[1688495016] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.85 ms
[1688495022] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688495106] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688495166] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 0.89 ms
[1688495196] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688495227] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;CRITICAL - Plugin timed out
[1688495286] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688495376] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688495441] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.87 ms
[1688495466] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688495520] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 75%, RTA = 0.92 ms
[1688495521] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.90 ms
[1688495664] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 79%, RTA = 1.05 ms
[1688495671] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688495761] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688495823] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688495874] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 93%, RTA = 1.01 ms
[1688495913] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688496003] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688496054] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688496089] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 1.22 ms
[1688496093] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688496164] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 68%, RTA = 0.97 ms
[1688496173] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 1.25 ms
[1688496263] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688496318] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 0.98 ms
[1688496353] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688496443] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688496444] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688496503] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 0.90 ms
[1688496533] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688496563] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688496598] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.92 ms
[1688496624] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688496673] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 68%, RTA = 1.06 ms
[1688496683] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.83 ms
[1688496713] HOST FLAPPING ALERT: SRVDNS;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1688496713] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688496803] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688496812] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.93 ms
[1688496822] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 54%, RTA = 1.02 ms
[1688496962] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 1.11 ms
[1688496962] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688497052] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688497112] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 0.85 ms
[1688497142] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688497232] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688497232] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688497322] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688497340] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 73%, RTA = 0.81 ms
[1688497370] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688497431] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 0.97 ms
[1688497472] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 0.71 ms
[1688497494] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688497562] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688497622] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 0.77 ms
[1688497652] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688497742] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688497776] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.98 ms
[1688497784] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 0.99 ms
[1688497922] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.80 ms
[1688497926] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688498012] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688498072] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 86%, RTA = 1.18 ms
[1688498102] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688498192] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688498282] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688498353] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 58%, RTA = 0.91 ms
[1688498372] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688498387] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1688498432] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 0.83 ms
[1688498442] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 1.01 ms
[1688498478] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 70%, RTA = 1.02 ms
[1688498507] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1688498612] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 5.57 ms
[1688498623] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688498702] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688498762] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.97 ms
[1688498792] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688498867] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 68%, RTA = 1.81 ms
[1688498882] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688498931] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.95 ms
[1688499060] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 1.13 ms
[1688499137] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688499157] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1688499210] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.86 ms
[1688499227] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688499277] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1688499317] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688499330] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688499407] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688499424] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1688499434] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 54%, RTA = 0.99 ms
[1688499574] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688499648] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688499690] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688499738] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688499828] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688499849] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 50%, RTA = 0.91 ms
[1688499857] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 44%, RTA = 1.42 ms
[1688499999] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688500061] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688500110] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 1.18 ms
[1688500151] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688500230] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688500241] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688500264] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.18 ms
[1688500282] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 73%, RTA = 0.79 ms
[1688500474] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688500531] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 93%, RTA = 0.96 ms
[1688500564] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688500590] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688500655] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688500685] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 1.10 ms
[1688500700] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 68%, RTA = 0.99 ms
[1688500775] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688500830] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 0.92 ms
[1688500865] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688500955] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688501045] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688501130] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688501135] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688501190] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 0.68 ms
[1688501225] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688501225] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.84 ms
[1688501235] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 54%, RTA = 0.93 ms
[1688501375] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688501414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688501430] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 1.00 ms
[1688501465] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688501534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688501550] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688501555] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688501645] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688501670] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 96%, RTA = 1.21 ms
[1688501709] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1688501730] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688501735] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688501765] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.96 ms
[1688501781] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 70%, RTA = 0.73 ms
[1688501829] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 34.48KB (0.03%), Average OUT: 256.06KB (0.20%)Total RX: 264.79GBytes, Total TX: 92.09GBytes
[1688501855] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688501909] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 83%, RTA = 1.07 ms
[1688501945] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688502035] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688502125] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688502150] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688502187] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 37%, RTA = 0.69 ms
[1688502215] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688502271] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.69 ms
[1688502337] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688502390] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688502422] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.70 ms
[1688502450] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.75 ms
[1688502510] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 1.40 ms
[1688502540] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688502604] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1688502660] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 90%, RTA = 0.72 ms
[1688502726] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 0.71 ms
[1688502750] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688502805] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 70%, RTA = 1.34 ms
[1688502879] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688502930] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688502969] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688503060] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 31.01 seconds)
[1688503150] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688503150] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 54%, RTA = 0.69 ms
[1688503160] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 54%, RTA = 0.75 ms
[1688503240] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688503299] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 0.89 ms
[1688503350] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 0.82 ms
[1688503360] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688503440] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688503500] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 0.97 ms
[1688503530] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688503620] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688503620] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688503660] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 54%, RTA = 0.88 ms
[1688503687] HOST ALERT: SRVDNS;UP;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 0.71 ms
[1688503750] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688503800] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 1.15 ms
[1688503840] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688503860] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688503930] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688504017] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 37%, RTA = 0.99 ms
[1688504020] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688504100] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 86%, RTA = 1.01 ms
[1688504107] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688504153] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 79%, RTA = 1.04 ms
[1688504183] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688504243] HOST ALERT: SRVDNS;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1688504280] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688504333] HOST ALERT: SRVDNS;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1688504400] HOST FLAPPING ALERT: SRVDNS;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1688504423] HOST ALERT: SRVDNS;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1688504434] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.93 ms
[1688504457] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 79%, RTA = 1.04 ms
[1688504524] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688504596] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688504640] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688504686] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688504700] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 0.98 ms
[1688504760] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688504776] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688504866] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688504931] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 77%, RTA = 0.84 ms
[1688504956] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688505022] HOST ALERT: SRVDNS;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1688505060] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 0.89 ms
[1688505083] HOST ALERT: SRVDNS;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1688505120] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688505173] HOST ALERT: SRVDNS;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1688505263] HOST ALERT: SRVDNS;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1688505345] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 68%, RTA = 0.62 ms
[1688505480] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 93%, RTA = 0.93 ms
[1688505688] HOST ALERT: SRVDNS;UP;HARD;1;PING OK - Packet loss = 16%, RTA = 0.60 ms
[1688505720] SERVICE ALERT: SRVDNS;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1688505780] SERVICE ALERT: SRVDNS;Ping;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1688505913] SERVICE ALERT: SRVDNS;Ping;OK;HARD;3;PING OK - Packet loss = 64%, RTA = 0.89 ms
[1688505930] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688506010] HOST ALERT: SRVDNS;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1688506063] SERVICE ALERT: SRVDNS;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 93%, RTA = 0.79 ms
[1688506100] HOST ALERT: SRVDNS;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1688506190] HOST ALERT: SRVDNS;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1688506280] HOST ALERT: SRVDNS;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1688506303] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1688506370] HOST ALERT: SRVDNS;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1688506434] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1688506457] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.96 ms
[1688506814] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688506934] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688507427] SERVICE FLAPPING ALERT: SRVDNS;Ping;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1688507900] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688507917] HOST ALERT: SRVDNS;UP;SOFT;1;PING OK - Packet loss = 50%, RTA = 0.74 ms
[1688508043] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688508044] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688508051] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688508056] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688508170] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688508174] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688508184] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688508188] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688508325] SERVICE ALERT: SRVDNS;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.2.30)
[1688508331] HOST ALERT: SRVDNS;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.30)
[1688508391] HOST ALERT: SRVDNS;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.30)
[1688508447] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.2.30)
[1688508453] HOST ALERT: SRVDNS;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.30)
[1688508516] HOST ALERT: SRVDNS;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.30)
[1688508576] HOST ALERT: SRVDNS;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.30)
[1688508601] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688508638] HOST ALERT: SRVDNS;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.30)
[1688508698] HOST ALERT: SRVDNS;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.30)
[1688508733] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688508758] HOST ALERT: SRVDNS;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.30)
[1688508821] HOST ALERT: SRVDNS;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.2.30)
[1688508881] HOST ALERT: SRVDNS;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.30)
[1688509202] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688509335] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688509486] HOST ALERT: SRVDNS;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1688509529] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1688509841] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1688509879] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1688509961] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 27.29KB (0.02%), Average OUT: 538.44KB (0.43%)Total RX: 8.00GBytes, Total TX: 9.06GBytes
[1688509999] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 103.70KB (0.08%), Average OUT: 1.92MB (1.54%)Total RX: 265.19GBytes, Total TX: 96.57GBytes
[1688510414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688510534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688511191] HOST ALERT: SRVDNS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1688511235] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1688511260] HOST ALERT: SRVDNS;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.30)
[1688511321] HOST ALERT: SRVDNS;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.30)
[1688511381] HOST ALERT: SRVDNS;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.30)
[1688511441] HOST ALERT: SRVDNS;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.30)
[1688511504] HOST ALERT: SRVDNS;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.30)
[1688511564] HOST ALERT: SRVDNS;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.30)
[1688511628] HOST ALERT: SRVDNS;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.30)
[1688511688] HOST ALERT: SRVDNS;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.2.30)
[1688511753] HOST ALERT: SRVDNS;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.30)
[1688512486] HOST ALERT: SRVDNS;UP;HARD;1;PING OK - Packet loss = 64%, RTA = 1.07 ms
[1688512529] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1688512586] SERVICE ALERT: SRVDNS;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.2.30)
[1688512595] HOST ALERT: SRVDNS;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.30)
[1688512655] HOST ALERT: SRVDNS;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.30)
[1688512706] SERVICE ALERT: SRVDNS;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.2.30)
[1688512715] HOST ALERT: SRVDNS;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.30)
[1688512776] HOST ALERT: SRVDNS;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.30)
[1688512839] HOST ALERT: SRVDNS;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.30)
[1688512899] HOST ALERT: SRVDNS;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.30)
[1688512960] HOST ALERT: SRVDNS;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.30)
[1688513020] HOST ALERT: SRVDNS;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.30)
[1688513083] HOST ALERT: SRVDNS;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.2.30)
[1688513143] HOST ALERT: SRVDNS;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.30)
[1688513446] SERVICE ALERT: SRV12WEB_cass.cl;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.2.14)
[1688513450] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.14)
[1688513462] SERVICE ALERT: SRV_chat.cass.cl;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1688513462] SERVICE ALERT: SRV_chat.cass.cl;Usuarios Login;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688513462] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688513466] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513466] SERVICE ALERT: SRV_chat.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513466] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTP;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513466] SERVICE ALERT: SRV_chat.cass.cl;Servicio Nginx;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513466] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.92)
[1688513466] SERVICE ALERT: SRV_chat.cass.cl;CPU;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513469] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTPS;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513469] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513471] SERVICE ALERT: SRV_chat.cass.cl;Proceso Docker;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688513471] SERVICE ALERT: SRV_chat.cass.cl;Servicio Docker;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688513475] SERVICE ALERT: SRV_chat.cass.cl;Puerto Contenedor;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688513476] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513479] SERVICE ALERT: SRV_chat.cass.cl;Puerto Portainer;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513487] SERVICE ALERT: SRV_chat.cass.cl;Memoria;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513513] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.14)
[1688513529] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.92)
[1688513570] SERVICE ALERT: SRV12WEB_cass.cl;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.2.14)
[1688513576] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.14)
[1688513585] SERVICE ALERT: SRV_chat.cass.cl;Usuarios Login;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513585] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513585] SERVICE ALERT: SRV_chat.cass.cl;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.92)
[1688513588] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513588] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTP;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513588] SERVICE ALERT: SRV_chat.cass.cl;Servicio Nginx;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513588] SERVICE ALERT: SRV_chat.cass.cl;Procesos;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: No route to host
[1688513591] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.92)
[1688513639] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.14)
[1688513651] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.92)
[1688513702] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.14)
[1688513711] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.92)
[1688513765] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.14)
[1688513771] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.92)
[1688513828] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.14)
[1688513831] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.92)
[1688513891] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.14)
[1688513891] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.92)
[1688513911] SERVICE ALERT: SRVDNS;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 0.81 ms
[1688513915] HOST ALERT: SRVDNS;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1688513951] HOST ALERT: SRV_chat.cass.cl;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.92)
[1688513955] HOST ALERT: SRV12WEB_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1688513990] SERVICE ALERT: SRV12WEB_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.08 ms
[1688514011] HOST ALERT: SRV_chat.cass.cl;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.92)
[1688514934] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688515066] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688518225] SERVICE ALERT: SRV_chat.cass.cl;Memoria;OK;HARD;1;[MEMORY] Total: 7919 MB - Used: 304 MB - 3% [SWAP] Total: 2048 MB - Used: 0 MB - 0%
[1688518241] SERVICE ALERT: SRV_chat.cass.cl;Servicio Docker;OK;HARD;1;OK: Service docker is running!
[1688518241] SERVICE ALERT: SRV_chat.cass.cl;Proceso Docker;OK;HARD;1;PROCS OK: 0 processes with command name 'docker'
[1688518245] HOST ALERT: SRV_chat.cass.cl;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1688518245] SERVICE ALERT: SRV_chat.cass.cl;Puerto Contenedor;OK;HARD;1;TCP OK - 0.000 second response time on port 8000
[1688518252] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTPS;OK;HARD;1;TCP OK - 0.000 second response time on port 443
[1688518255] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;SOFT;1;DISK UNKNOWN - free space:
[1688518262] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1688518263] SERVICE ALERT: SRV_chat.cass.cl;Usuarios Login;OK;HARD;1;USERS OK - 0 users currently logged in
[1688518266] SERVICE ALERT: SRV_chat.cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.07 ms
[1688518266] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTP;OK;HARD;1;TCP OK - 0.000 second response time on port 80
[1688518266] SERVICE ALERT: SRV_chat.cass.cl;Servicio Nginx;OK;HARD;1;OK: Service nginx is running!
[1688518267] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 314 processes
[1688518267] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;OK;HARD;1;PROCS OK: 3 processes with command name 'nginx'
[1688518278] SERVICE ALERT: SRV_chat.cass.cl;Puerto Portainer;OK;HARD;1;TCP OK - 0.000 second response time on port 9443
[1688518288] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1688518314] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;SOFT;2;DISK UNKNOWN - free space:
[1688518319] SERVICE ALERT: SRV_chat.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 1.19, 0.66, 0.26
[1688518418] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1688518434] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;HARD;3;DISK UNKNOWN - free space:
[1688518550] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688518801] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688518854] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;OK;HARD;1;DISK OK - free space: /var/tmp 972536 MiB (93.05% inode=-):
[1688518934] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1688519414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1688519534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!
[1688525402] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1688525402] SERVICE ALERT: SRV_chat.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1688525522] SERVICE ALERT: SRV_chat.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 0 users currently logged in
[1688525522] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1688528414] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1688528534] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;SOFT;2;OK: Service vibe is running!