[1681886941] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1681886943] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 2.90 ms
[1681886968] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.63 ms
[1681887067] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1681898621] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1681898621] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1681898741] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 27 day(s) 17 hour(s) 49 minute(s)
[1681898741] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.01, 0.01, 0.00
[1681905889] SERVICE ALERT: Veeam2023.cass.pmc;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1681905889] SERVICE ALERT: Veeam2023.cass.pmc;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1681905895] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1681905910] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.14)
[1681905910] SERVICE ALERT: SRV12WEB_cass.cl;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.2.14)
[1681905911] SERVICE ALERT: Veeam2023.cass.pmc;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.83)
[1681905911] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.83)
[1681905925] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;HARD;1;connect to address 10.0.0.83 and port 12489: No route to host
[1681905973] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.14)
[1681905974] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.83)
[1681906012] SERVICE ALERT: Veeam2023.cass.pmc;Memoria;CRITICAL;HARD;3;connect to address 10.0.0.83 and port 12489: No route to host
[1681906012] SERVICE ALERT: Veeam2023.cass.pmc;Tiempo Activo;CRITICAL;HARD;3;connect to address 10.0.0.83 and port 12489: No route to host
[1681906018] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;CRITICAL;HARD;3;connect to address 10.0.0.83 and port 12489: No route to host
[1681906034] SERVICE ALERT: Veeam2023.cass.pmc;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.83)
[1681906036] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.14)
[1681906037] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.83)
[1681906097] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.83)
[1681906099] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.14)
[1681906157] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.83)
[1681906162] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.14)
[1681906220] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.83)
[1681906225] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.14)
[1681906283] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.83)
[1681906288] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.14)
[1681906345] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.83)
[1681906351] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.14)
[1681906405] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.83)
[1681906412] HOST ALERT: SRV12WEB_cass.cl;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.2.14)
[1681906421] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1681906465] HOST ALERT: Veeam2023.cass.pmc;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.83)
[1681906475] HOST ALERT: SRV12WEB_cass.cl;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.14)
[1681906541] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 27 day(s) 19 hour(s) 59 minute(s)
[1681906931] SERVICE ALERT: SRV11VIBE_cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906931] SERVICE ALERT: SRV11VIBE_cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906931] SERVICE ALERT: SRV11VIBE_cass.cl;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.2.13)
[1681906934] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.13)
[1681906948] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906953] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906953] SERVICE ALERT: SRV11VIBE_cass.cl;Almacenamiento /;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906953] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906963] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906963] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906966] SERVICE ALERT: SRV11VIBE_cass.cl;Procesos Zombies;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681906990] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.2.13)
[1681907050] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.2.13)
[1681907053] SERVICE ALERT: SRV11VIBE_cass.cl;Usuarios Login;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681907053] SERVICE ALERT: SRV11VIBE_cass.cl;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.2.13)
[1681907053] SERVICE ALERT: SRV11VIBE_cass.cl;Procesos;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: No route to host
[1681907110] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.2.13)
[1681907170] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.2.13)
[1681907177] SERVICE ALERT: SRVAD2.cass.pmc;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.5)
[1681907180] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.5)
[1681907196] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.11)
[1681907204] SERVICE ALERT: SRV05AD02_cass.cl;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.0.11)
[1681907220] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.123)
[1681907223] SERVICE ALERT: SRVDC1.cass.pmc;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.0.123)
[1681907230] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.2.13)
[1681907243] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.5)
[1681907259] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.11)
[1681907280] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.120)
[1681907282] HOST ALERT: SRVMFSD;UNREACHABLE;HARD;10;CRITICAL - Host Unreachable (10.0.2.12)
[1681907283] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.123)
[1681907288] SERVICE ALERT: Host-10.0.0.120;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.0.120)
[1681907290] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.2.13)
[1681907300] SERVICE ALERT: SRVAD2.cass.pmc;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.5)
[1681907306] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.5)
[1681907322] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.11)
[1681907343] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.120)
[1681907343] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.123)
[1681907350] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.2.13)
[1681907369] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.5)
[1681907382] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 3 status: 2)
[1681907384] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.11)
[1681907389] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;SOFT;1;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681907403] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.123)
[1681907406] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.120)
[1681907410] HOST ALERT: SRV11VIBE_cass.cl;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.2.13)
[1681907432] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.5)
[1681907447] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.11)
[1681907463] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.123)
[1681907468] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.120)
[1681907470] HOST ALERT: SRV11VIBE_cass.cl;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.13)
[1681907495] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.5)
[1681907503] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;2;CRITICAL: Interface is 'down' (interface: 3 status: 2)
[1681907509] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;SOFT;2;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681907510] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.11)
[1681907523] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.123)
[1681907528] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.120)
[1681907558] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.5)
[1681907573] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.11)
[1681907583] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.123)
[1681907588] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.120)
[1681907621] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.5)
[1681907623] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 0.00B (0.00%), Average OUT: 62.47B (0.00%)Total RX: 7222.88GBytes, Total TX: 19211.40GBytes
[1681907629] SERVICE ALERT: Servidores CASS;Estado Puerto 3;OK;SOFT;3;OK: Interface GigabitEthernet1/0/3 (index 3) is up.
[1681907636] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.11)
[1681907643] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.123)
[1681907651] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.120)
[1681907684] HOST ALERT: SRVAD2.cass.pmc;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.5)
[1681907699] HOST ALERT: SRV05AD02_cass.cl;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.11)
[1681907704] HOST ALERT: SRVDC1.cass.pmc;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.123)
[1681907714] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.120)
[1681907747] HOST ALERT: SRVAD2.cass.pmc;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.5)
[1681907762] HOST ALERT: SRV05AD02_cass.cl;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.11)
[1681907767] HOST ALERT: SRVDC1.cass.pmc;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.123)
[1681907777] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.120)
[1681907840] HOST ALERT: Host-10.0.0.120;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.120)
[1681907870] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;SOFT;1;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681907889] SERVICE ALERT: Host-10.0.0.120;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1681907893] HOST ALERT: Host-10.0.0.120;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1681907943] HOST ALERT: SRVMFSD;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.12)
[1681907990] SERVICE ALERT: Servidores CASS;Estado Puerto 3;OK;SOFT;2;OK: Interface GigabitEthernet1/0/3 (index 3) is up.
[1681908076] SERVICE ALERT: Host-10.0.0.120;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.120)
[1681908076] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.120)
[1681908103] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;1;CRITICAL: Interface is 'down' (interface: 3 status: 2)
[1681908110] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;SOFT;1;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681908126] HOST ALERT: SRVMFSD;UNREACHABLE;HARD;10;CRITICAL - Host Unreachable (10.0.2.12)
[1681908139] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.120)
[1681908199] SERVICE ALERT: Host-10.0.0.120;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.120)
[1681908202] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.120)
[1681908224] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;2;CRITICAL: Interface is 'down' (interface: 3 status: 2)
[1681908230] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;SOFT;2;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681908265] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.120)
[1681908328] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.120)
[1681908344] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL: Interface is 'down' (interface: 3 status: 2)
[1681908350] SERVICE ALERT: Servidores CASS;Estado Puerto 3;CRITICAL;HARD;3;CRITICAL: Interface GigabitEthernet1/0/3 (index 3) is down.
[1681908391] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.120)
[1681908404] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 0.00B (0.00%), Average OUT: 162.07B (0.00%)Total RX: 7222.90GBytes, Total TX: 19211.40GBytes
[1681908410] SERVICE ALERT: Servidores CASS;Estado Puerto 3;OK;HARD;1;OK: Interface GigabitEthernet1/0/3 (index 3) is up.
[1681908454] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.120)
[1681908517] HOST ALERT: Host-10.0.0.120;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.120)
[1681908581] HOST ALERT: Host-10.0.0.120;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1681908584] HOST ALERT: SRVMFSD;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.2.12)
[1681908620] SERVICE ALERT: Host-10.0.0.120;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1681908711] HOST ALERT: SRVAD2.cass.pmc;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1681908741] SERVICE ALERT: SRVAD2.cass.pmc;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1681908791] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;HARD;1;[MEMORY] Total: 7693 MB - Used: 362 MB - 4% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1681908791] HOST ALERT: SRV11VIBE_cass.cl;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.90 ms
[1681908792] SERVICE ALERT: SRV11VIBE_cass.cl;Usuarios Login;OK;HARD;1;USERS OK - 0 users currently logged in
[1681908792] SERVICE ALERT: SRV11VIBE_cass.cl;Procesos;OK;HARD;1;PROCS OK: 261 processes
[1681908795] SERVICE ALERT: SRV11VIBE_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.00 ms
[1681908810] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;HARD;1;PROCS OK: 0 processes with command name 'vibe'
[1681908811] SERVICE ALERT: SRV11VIBE_cass.cl;Servicio Vibe;OK;HARD;1;OK: Service vibe is running!
[1681908811] SERVICE ALERT: SRV11VIBE_cass.cl;Almacenamiento /;OK;HARD;1;DISK OK - free space: /var/tmp 78775 MiB (86.00% inode=-):
[1681908820] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.35, 0.10, 0.04
[1681908821] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;OK;HARD;1;Uptime OK: 0 day(s) 0 hour(s) 0 minute(s)
[1681908824] SERVICE ALERT: SRV11VIBE_cass.cl;Procesos Zombies;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1681908843] SERVICE ALERT: SRV12WEB_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1681908847] HOST ALERT: SRV12WEB_cass.cl;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1681908895] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;OK;HARD;1;c:\ - total: 99.33 Gb - used: 81.30 Gb (82%) - free 18.03 Gb (18%)
[1681908899] HOST ALERT: Veeam2023.cass.pmc;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1681908915] SERVICE ALERT: Veeam2023.cass.pmc;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.31 ms
[1681908922] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 0% (5 min average)
[1681908949] SERVICE ALERT: Veeam2023.cass.pmc;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1681908949] SERVICE ALERT: Veeam2023.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:18815.17 MB - used: 1785.64 MB (9%) - free: 17029.54 MB (91%)
[1681909117] SERVICE ALERT: SRV_SUM;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 83%, RTA = 0.55 ms
[1681909117] HOST ALERT: SRV_SUM;UP;HARD;1;PING OK - Packet loss = 44%, RTA = 0.62 ms
[1681909152] SERVICE ALERT: SRV_SUM;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1681909164] HOST ALERT: SRV05AD02_cass.cl;UP;HARD;1;PING OK - Packet loss = 78%, RTA = 1.59 ms
[1681909185] SERVICE ALERT: SRV05AD02_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1681909243] HOST ALERT: SRV_sync.cass.cl;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.03 ms
[1681909279] SERVICE ALERT: SRV_sync.cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.85 ms
[1681909324] SERVICE ALERT: SRVDC1.cass.pmc;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1681909328] HOST ALERT: SRVDC1.cass.pmc;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.00 ms
[1681909430] HOST ALERT: SRVMFSD;UP;HARD;1;PING OK - Packet loss = 28%, RTA = 1.20 ms
[1681909444] SERVICE ALERT: SRVMFSD;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.96 ms
[1681909448] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.2.6)
[1681909451] HOST ALERT: SRV_Oraculo_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.6)
[1681909515] HOST ALERT: SRV_Oraculo_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.84 ms
[1681909543] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681909572] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.11 ms
[1681909676] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681909801] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681909933] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681913197] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1681913317] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 23.27KB (0.02%), Average OUT: 20.06KB (0.02%)Total RX: 16.96GBytes, Total TX: 7.71GBytes
[1681915532] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1681915541] SERVICE ALERT: Veeam2023.cass.pmc;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 1.50 ms
[1681915544] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.83)
[1681915552] SERVICE ALERT: Veeam2023.cass.pmc;Memoria;CRITICAL;HARD;1;connect to address 10.0.0.83 and port 12489: No route to host
[1681915552] SERVICE ALERT: Veeam2023.cass.pmc;Tiempo Activo;CRITICAL;HARD;1;connect to address 10.0.0.83 and port 12489: No route to host
[1681915558] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;CRITICAL;HARD;1;connect to address 10.0.0.83 and port 12489: No route to host
[1681915604] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.83)
[1681915655] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;HARD;3;connect to address 10.0.0.83 and port 12489: No route to host
[1681915664] SERVICE ALERT: Veeam2023.cass.pmc;Ping;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.83)
[1681915667] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.83)
[1681915721] SERVICE ALERT: Pasmar;Ping;CRITICAL;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.97)
[1681915721] HOST ALERT: Pasmar;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.97)
[1681915721] HOST ALERT: FW Pasmar;UNREACHABLE;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.98)
[1681915721] SERVICE ALERT: FW Pasmar;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.98)
[1681915730] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.83)
[1681915785] HOST ALERT: Pasmar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1681915792] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.0.0.83)
[1681915811] HOST ALERT: FW Pasmar;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1681915845] SERVICE ALERT: Pasmar;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.44 ms
[1681915855] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.0.0.83)
[1681915875] HOST ALERT: FW Pasmar;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1681915905] SERVICE ALERT: FW Pasmar;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.15 ms
[1681915909] HOST ALERT: FW Pasmar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.19 ms
[1681915915] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.0.0.83)
[1681915975] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.0.0.83)
[1681916035] HOST ALERT: Veeam2023.cass.pmc;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.0.0.83)
[1681916095] HOST ALERT: Veeam2023.cass.pmc;DOWN;HARD;10;CRITICAL - Host Unreachable (10.0.0.83)
[1681916493] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 0% (5 min average)
[1681916497] HOST ALERT: Veeam2023.cass.pmc;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.47 ms
[1681916505] SERVICE ALERT: Veeam2023.cass.pmc;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1681916509] SERVICE ALERT: Veeam2023.cass.pmc;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1681916509] SERVICE ALERT: Veeam2023.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:18815.17 MB - used: 1760.35 MB (9%) - free: 17054.82 MB (91%)
[1681916515] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;OK;HARD;1;c:\ - total: 99.33 Gb - used: 81.14 Gb (82%) - free 18.19 Gb (18%)
[1681921801] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681921921] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1681925246] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1681925367] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1681926051] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1681926171] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 29.75B (0.00%), Average OUT: 8.39B (0.00%)Total RX: 1816.95GBytes, Total TX: 3964.42GBytes
[1681926356] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1681926362] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1681926481] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1681926492] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1681926504] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1681926577] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1681926579] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681926612] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 32.46B (0.00%), Average OUT: 7.09B (0.00%)Total RX: 1816.95GBytes, Total TX: 3964.42GBytes
[1681926624] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.026 sec. response time
[1681926641] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1681926697] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1681927109] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;WARNING;SOFT;1;WARNING - IN bandwidth (88.08%) too high
[1681927229] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 31.61MB (25.29%), Average OUT: 365.26KB (0.29%)Total RX: 341.16GBytes, Total TX: 319.09GBytes
[1681929001] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1681929121] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1681930963] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681931095] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681931143] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681931276] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681931401] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681931402] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681931534] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681931534] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681934564] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681934564] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681934583] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681934696] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681934697] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681934715] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681935001] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681935002] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681935133] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681935134] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681938279] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (210819928393.28%) too high
[1681938400] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;2;CRITICAL - IN bandwidth (2243.37%) too high
[1681938520] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;3;OK - Average IN: 46.39KB (0.04%), Average OUT: 636.53KB (0.51%)Total RX: 341.83GBytes, Total TX: 324.62GBytes
[1681939801] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1681939921] SERVICE ALERT: SRV_chat.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1681954202] HOST ALERT: FW Pasmar;DOWN;SOFT;1;CRITICAL - Time to live exceeded (190.121.2.98)
[1681954240] SERVICE ALERT: FW Pasmar;Ping;CRITICAL;HARD;1;CRITICAL - Time to live exceeded (190.121.2.98)
[1681954292] HOST ALERT: FW Pasmar;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1681954382] HOST ALERT: FW Pasmar;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1681954387] SERVICE ALERT: FW Pasmar;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 82%, RTA = 1.29 ms
[1681954424] SERVICE ALERT: FW Pasmar;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.17 ms
[1681954428] HOST ALERT: FW Pasmar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.42 ms
[1681955554] HOST ALERT: FW Clinica PV Telsur;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1681955554] HOST ALERT: FW Clinica PV Movistar;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1681955560] SERVICE ALERT: FW Clinica PV Movistar;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1681955644] HOST ALERT: FW Clinica PV Telsur;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1681955644] HOST ALERT: FW Clinica PV Movistar;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1681955734] HOST ALERT: FW Clinica PV Telsur;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1681955734] HOST ALERT: FW Clinica PV Movistar;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1681955774] SERVICE ALERT: FW Clinica PV Movistar;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 34.45 ms
[1681955778] HOST ALERT: FW Clinica PV Movistar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 34.31 ms
[1681955798] HOST ALERT: FW Clinica PV Telsur;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.81 ms