[1675484988] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675485121] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675487084] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.37, 0.13, 0.05
[1675487204] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.18, 0.14, 0.07
[1675488251] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;The command (check_updates) returned an invalid return code: 255
[1675488372] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1675491584] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.14, 0.14, 0.10
[1675491704] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.13, 0.14, 0.11
[1675491824] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1675495992] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;The command (check_updates) returned an invalid return code: 255
[1675496113] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1675506704] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.41, 0.22, 0.17
[1675507184] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.02, 0.18, 0.19
[1675507904] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.01, 0.03, 0.10
[1675515212] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675515217] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675515227] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675515247] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675515344] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675515349] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675515358] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675515380] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675541750] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675541871] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1675549167] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1675549290] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1675549412] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional