[1690689738] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.17%) too high
[1690689858] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (95.11%) too high
[1690689978] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (96.48%) too high
[1690690519] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 799.12KB (0.64%), Average OUT: 101.48MB (81.19%)Total RX: 32376.08GBytes, Total TX: 75904.09GBytes
[1690690578] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.61%) too high
[1690690698] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 3.14MB (2.51%), Average OUT: 103.50MB (82.80%)Total RX: 32376.50GBytes, Total TX: 75923.05GBytes
[1690690819] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.38%) too high
[1690690939] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.18%) too high
[1690691059] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (90.59%) too high
[1690691179] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1690691179] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 520.72KB (0.42%), Average OUT: 103.04MB (82.43%)Total RX: 32377.08GBytes, Total TX: 75975.98GBytes
[1690691300] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.65%) too high
[1690691420] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 513.53KB (0.41%), Average OUT: 104.79MB (83.83%)Total RX: 32377.20GBytes, Total TX: 76001.61GBytes
[1690691601] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (90.55%) too high
[1690691721] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (93.03%) too high
[1690691841] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (90.90%) too high
[1690692021] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 508.99KB (0.41%), Average OUT: 104.99MB (83.99%)Total RX: 32377.53GBytes, Total TX: 76068.35GBytes
[1690692082] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.36%) too high
[1690692202] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.92%) too high
[1690692322] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (85.30%) too high
[1690692442] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 478.24KB (0.38%), Average OUT: 96.26MB (77.01%)Total RX: 32377.75GBytes, Total TX: 76114.10GBytes
[1690692802] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.49%) too high
[1690692922] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (85.99%) too high
[1690693043] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 512.14KB (0.41%), Average OUT: 105.00MB (84.00%)Total RX: 32378.05GBytes, Total TX: 76174.86GBytes
[1690693163] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.19%) too high
[1690693283] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 487.79KB (0.39%), Average OUT: 100.99MB (80.79%)Total RX: 32378.18GBytes, Total TX: 76199.76GBytes
[1690693463] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (94.16%) too high
[1690693583] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 524.73KB (0.42%), Average OUT: 103.81MB (83.04%)Total RX: 32378.33GBytes, Total TX: 76231.42GBytes
[1690693704] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.61%) too high
[1690693824] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.95%) too high
[1690693944] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (86.39%) too high
[1690694124] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 496.41KB (0.40%), Average OUT: 99.28MB (79.42%)Total RX: 32378.63GBytes, Total TX: 76291.63GBytes
[1690694185] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.97%) too high
[1690694305] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (94.62%) too high
[1690694425] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (91.50%) too high
[1690694545] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 518.63KB (0.41%), Average OUT: 103.94MB (83.15%)Total RX: 32378.87GBytes, Total TX: 76339.23GBytes
[1690694605] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (93.83%) too high
[1690694726] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.98%) too high
[1690694846] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (87.68%) too high
[1690694967] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 511.90KB (0.41%), Average OUT: 103.49MB (82.80%)Total RX: 32379.11GBytes, Total TX: 76386.55GBytes
[1690695026] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (91.30%) too high
[1690695146] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (88.77%) too high
[1690695267] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (94.39%) too high
[1690695567] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 516.00KB (0.41%), Average OUT: 102.63MB (82.11%)Total RX: 32379.45GBytes, Total TX: 76454.51GBytes
[1690695628] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.15%) too high
[1690695634] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690695748] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 489.74KB (0.39%), Average OUT: 100.02MB (80.01%)Total RX: 32379.54GBytes, Total TX: 76473.23GBytes
[1690695754] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4595 MB - 59% [SWAP] Total: 7944 MB - Used: 96 MB - 1%
[1690695929] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.13%) too high
[1690696049] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (85.56%) too high
[1690696169] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 512.20KB (0.41%), Average OUT: 102.81MB (82.25%)Total RX: 32379.76GBytes, Total TX: 76517.24GBytes
[1690696290] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.83%) too high
[1690696410] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.35%) too high
[1690696530] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 539.61KB (0.43%), Average OUT: 80.53MB (64.42%)Total RX: 32379.99GBytes, Total TX: 76553.44GBytes
[1690697250] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1690699234] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690699354] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4591 MB - 59% [SWAP] Total: 7944 MB - Used: 96 MB - 1%
[1690699801] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690699921] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.006 sec. response time
[1690700459] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1690700461] SERVICE ALERT: FW Gore Los Lagos;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1690700465] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1690700471] SERVICE ALERT: Gore Los Lagos;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1690700495] HOST ALERT: FW Gore Los Lagos;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1690700500] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 50%, RTA = 1.22 ms
[1690700504] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.14 ms
[1690700504] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1690700508] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.20 ms
[1690705801] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25: Connection reset by peer
[1690705921] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.025 sec. response time
[1690708234] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690708354] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4586 MB - 59% [SWAP] Total: 7944 MB - Used: 96 MB - 1%
[1690710083] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;1;YUM CRITICAL: Yum nagios plugin has self terminated after exceeding the timeout (30 seconds)
[1690710145] SERVICE ALERT: SRV_nagios;Actualizacion;OK;SOFT;2;YUM OK: 0 Updates Available
[1690711201] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1690711321] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.023 sec. response time
[1690712401] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690712524] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 3.035 sec. response time
[1690718671] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (86.33%) too high
[1690718791] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - IN bandwidth (88.04%) too high
[1690718911] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (85.61%) too high
[1690718972] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 104.87MB (83.89%), Average OUT: 47.84MB (38.27%)Total RX: 32426.60GBytes, Total TX: 76968.22GBytes
[1690719034] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1690719154] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4586 MB - 59% [SWAP] Total: 7944 MB - Used: 101 MB - 1%
[1690724434] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1690724555] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4586 MB - 59% [SWAP] Total: 7944 MB - Used: 101 MB - 1%
[1690725811] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (87.31%) too high
[1690725932] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 66.96MB (53.57%), Average OUT: 31.93MB (25.54%)Total RX: 32503.29GBytes, Total TX: 77063.41GBytes
[1690730024] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690730086] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1690736792] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (89.01%) too high
[1690736912] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 101.61MB (81.29%), Average OUT: 17.05MB (13.64%)Total RX: 32571.18GBytes, Total TX: 77145.30GBytes
[1690737092] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (92.17%) too high
[1690737212] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - IN bandwidth (85.18%) too high
[1690737333] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (85.75%) too high
[1690737453] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 101.50MB (81.20%), Average OUT: 21.18MB (16.94%)Total RX: 32628.57GBytes, Total TX: 77155.26GBytes
[1690737573] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (92.25%) too high
[1690737693] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - IN bandwidth (91.50%) too high
[1690737814] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (87.59%) too high
[1690737874] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1690737874] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 104.35MB (83.48%), Average OUT: 19.90MB (15.92%)Total RX: 32674.92GBytes, Total TX: 77162.02GBytes
[1690738294] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (88.83%) too high
[1690738415] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 104.23MB (83.38%), Average OUT: 14.93MB (11.94%)Total RX: 32729.78GBytes, Total TX: 77172.04GBytes
[1690738535] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (86.19%) too high
[1690738655] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 99.33MB (79.46%), Average OUT: 18.85MB (15.08%)Total RX: 32754.62GBytes, Total TX: 77176.53GBytes
[1690739075] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (87.24%) too high
[1690739195] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 99.67MB (79.74%), Average OUT: 19.61MB (15.69%)Total RX: 32809.38GBytes, Total TX: 77187.33GBytes
[1690739375] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (85.71%) too high
[1690739402] SERVICE ALERT: SRV_SOS2021.cass.cl;Almacenamiento /;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1690739495] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 101.72MB (81.38%), Average OUT: 19.92MB (15.94%)Total RX: 32839.89GBytes, Total TX: 77193.58GBytes
[1690739522] SERVICE ALERT: SRV_SOS2021.cass.cl;Almacenamiento /;OK;SOFT;2;DISK OK - free space: /var/tmp 118412 MiB (91.21% inode=-):
[1690739735] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (87.09%) too high
[1690739856] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1690739856] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 103.89MB (83.11%), Average OUT: 20.09MB (16.07%)Total RX: 32877.35GBytes, Total TX: 77200.59GBytes
[1690739976] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (87.74%) too high
[1690740096] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 100.94MB (80.75%), Average OUT: 20.03MB (16.02%)Total RX: 32902.62GBytes, Total TX: 77205.07GBytes
[1690740336] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (88.07%) too high
[1690740457] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 40.58MB (32.47%), Average OUT: 8.53MB (6.82%)Total RX: 32932.98GBytes, Total TX: 77210.35GBytes
[1690742437] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (95.41%) too high
[1690742558] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - IN bandwidth (94.06%) too high
[1690742678] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (96.31%) too high
[1690742858] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 90.89MB (72.71%), Average OUT: 20.51MB (16.41%)Total RX: 33049.97GBytes, Total TX: 77242.42GBytes
[1690742979] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (89.14%) too high
[1690743099] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - IN bandwidth (91.23%) too high
[1690743219] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (88.25%) too high
[1690744239] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - IN bandwidth (100.14%) too high
[1690744299] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (95.54%) too high
[1690744839] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - IN bandwidth (98.14%) too high
[1690744900] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1690744900] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - IN bandwidth (91.87%) too high
[1690745261] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 101.48MB (81.19%), Average OUT: 21.66MB (17.33%)Total RX: 33328.73GBytes, Total TX: 77277.33GBytes
[1690746006] SERVICE ALERT: Qnap NAS;Sysinfo;CRITICAL;SOFT;1;CRITICAL: SNMP to 10.0.0.28 is not available or wrong community string
[1690746100] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1690746126] SERVICE ALERT: Qnap NAS;Sysinfo;OK;SOFT;2;NAS Cloud, Model TS-879U-RP, Firmware 4.3.6, Max HD number 8, No. Volume 1
[1690746400] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - IN bandwidth (92.45%) too high
[1690746520] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 88.76MB (71.01%), Average OUT: 24.89MB (19.91%)Total RX: 33413.23GBytes, Total TX: 77303.35GBytes
[1690747201] SERVICE ALERT: SRV_chat.cass.cl;Procesos;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690747322] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 303 processes
[1690758034] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1690758155] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4584 MB - 59% [SWAP] Total: 7944 MB - Used: 101 MB - 1%
[1690763435] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1690763555] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4595 MB - 59% [SWAP] Total: 7944 MB - Used: 101 MB - 1%
[1690764374] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1690764506] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690769980] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.72%) too high
[1690770101] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 16.51MB (13.21%), Average OUT: 96.69MB (77.35%)Total RX: 34248.13GBytes, Total TX: 78278.43GBytes
[1690772621] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (93.47%) too high
[1690772742] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.56%) too high
[1690772862] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (91.64%) too high
[1690773042] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.95%) too high
[1690773102] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (97.84%) too high
[1690773222] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1690773222] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (99.67%) too high
[1690773343] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.18%) too high
[1690773403] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 2.21MB (1.76%), Average OUT: 50.74MB (40.60%)Total RX: 34284.07GBytes, Total TX: 78547.87GBytes
[1690774235] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1690774242] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1690774256] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1690774303] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1690774303] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1690774356] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4591 MB - 59% [SWAP] Total: 7944 MB - Used: 101 MB - 1%
[1690774374] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774389] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774436] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774436] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774543] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1690774801] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690774801] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690774802] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1690774927] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774933] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1690774934] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available