[1702184401] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1702184521] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 7 day(s) 21 hour(s) 37 minute(s)
[1702188046] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1702188046] SERVICE ALERT: FW Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 1.24 ms
[1702188051] HOST ALERT: FW Gore Los Lagos;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1702188081] HOST ALERT: Gore Los Lagos;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1702188090] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 1.21 ms
[1702188110] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.20 ms
[1702188124] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.05 ms
[1702188128] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.14 ms
[1702188170] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1702197001] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1702197121] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 8 day(s) 1 hour(s) 7 minute(s)
[1702214401] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.0.20: Connection reset by peer
[1702214521] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 8 day(s) 5 hour(s) 57 minute(s)
[1702215974] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1702216098] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1702217603] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1702217679] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1702219801] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.0.20: Connection reset by peer
[1702219921] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 8 day(s) 7 hour(s) 27 minute(s)
[1702225801] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1702225921] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 8 day(s) 9 hour(s) 7 minute(s)
[1702241465] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702242005] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702242125] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702242245] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702242965] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702243445] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702243565] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702243685] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702244705] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702244945] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702245066] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702245186] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702247046] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702247346] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702247466] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702247586] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1702248481] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1702248543] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1702256750] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.02%) too high
[1702256870] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 43.31MB (34.65%), Average OUT: 69.00MB (55.20%)Total RX: 1357.78GBytes, Total TX: 14669.16GBytes