[1673672625] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.25, 0.08, 0.03
[1673672745] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;SOFT;2;OK - load average per CPU: 0.05, 0.06, 0.03
[1673677425] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.13, 0.13, 0.10
[1673677545] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.13, 0.13, 0.11
[1673677665] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.12, 0.13, 0.11
[1673684040] SERVICE ALERT: Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 19.24 ms
[1673684164] SERVICE ALERT: Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 17.97 ms
[1673692245] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.44, 0.23, 0.16
[1673692365] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.17, 0.22, 0.17
[1673692845] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;HARD;1;OK - load average per CPU: 0.00, 0.04, 0.10
[1673718089] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 70%, RTA = 36.52 ms
[1673718095] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1673718107] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1673718125] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1673718138] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1673718140] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1673718168] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1673718215] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1673718219] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1673718225] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1673718228] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1673718258] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1673718305] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1673718318] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1673718348] HOST ALERT: ADL Diagnostic;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1673718395] HOST ALERT: FW Cass Entel;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1673718408] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1673718438] HOST ALERT: ADL Diagnostic;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1673718485] HOST ALERT: FW Cass Entel;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1673718498] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1673718528] HOST ALERT: ADL Diagnostic;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1673718575] HOST ALERT: FW Cass Entel;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1673718588] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1673718618] HOST ALERT: ADL Diagnostic;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1673718665] HOST ALERT: FW Cass Entel;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1673718678] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1673718708] HOST ALERT: ADL Diagnostic;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1673718755] HOST ALERT: FW Cass Entel;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1673718768] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1673718798] HOST ALERT: ADL Diagnostic;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1673718845] HOST ALERT: FW Cass Entel;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1673718858] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;9;(Host check timed out after 30.00 seconds)
[1673718888] HOST ALERT: ADL Diagnostic;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1673718935] HOST ALERT: FW Cass Entel;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1673718948] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1673719173] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 37.13 ms
[1673719177] HOST ALERT: ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 37.71 ms
[1673719180] HOST ALERT: FW Cass Entel;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 35.18 ms
[1673719180] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.34 ms
[1673719181] HOST ALERT: FW ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 35.09 ms
[1673719214] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.03 ms
[1673719301] SERVICE ALERT: FW Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 3.19 ms
[1673719425] SERVICE ALERT: FW Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.17 ms
[1673719659] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1673719663] SERVICE ALERT: FW Confe;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 37.89 ms
[1673719666] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1673719683] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1673719685] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1673719687] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1673719700] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1673719717] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1673719756] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1673719773] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1673719787] SERVICE ALERT: FW Confe;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 38.27 ms
[1673719789] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1673719796] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1673719807] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1673719846] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1673719864] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1673719897] HOST ALERT: ADL Diagnostic;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1673719910] SERVICE ALERT: SRV_Filr4_restore1123;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 1.21 ms
[1673719936] HOST ALERT: FW Cass Entel;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1673719954] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1673719987] HOST ALERT: ADL Diagnostic;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1673720026] HOST ALERT: FW Cass Entel;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1673720034] SERVICE ALERT: SRV_Filr4_restore1123;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.99 ms
[1673720044] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1673720077] HOST ALERT: ADL Diagnostic;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1673720116] HOST ALERT: FW Cass Entel;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1673720134] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1673720167] HOST ALERT: ADL Diagnostic;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1673720206] HOST ALERT: FW Cass Entel;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1673720224] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1673720257] HOST ALERT: ADL Diagnostic;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1673720296] HOST ALERT: FW Cass Entel;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1673720314] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;8;(Host check timed out after 30.00 seconds)
[1673720347] HOST ALERT: ADL Diagnostic;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1673720386] HOST ALERT: FW Cass Entel;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1673720393] SERVICE ALERT: FW Confe;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 37.81 ms
[1673720404] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1673720437] HOST ALERT: ADL Diagnostic;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1673720476] HOST ALERT: FW Cass Entel;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1673720494] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1673720517] SERVICE ALERT: FW Confe;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 38.14 ms
[1673720871] SERVICE ALERT: FW Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 3.13 ms
[1673720995] SERVICE ALERT: FW Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.28 ms
[1673721120] SERVICE ALERT: SRV_Filr4_restore1123;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 1.29 ms
[1673721173] SERVICE ALERT: SRVDNS;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 37%, RTA = 1.00 ms
[1673721244] SERVICE ALERT: SRV_Filr4_restore1123;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.14 ms
[1673721297] SERVICE ALERT: SRVDNS;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.05 ms
[1673721302] SERVICE ALERT: FW Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 37%, RTA = 3.21 ms
[1673721426] SERVICE ALERT: FW Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.23 ms
[1673722233] HOST ALERT: FW Cass Entel;UP;HARD;1;PING OK - Packet loss = 72%, RTA = 33.73 ms
[1673722239] HOST ALERT: FW ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 33.73 ms
[1673722241] HOST ALERT: ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 34.99 ms
[1673722243] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 34.93 ms
[1673722251] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 33.66 ms
[1673722274] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 33.81 ms
[1673745286] SERVICE ALERT: Veeam2019.cass.pmc;NsClient++ CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1673745406] SERVICE ALERT: Veeam2019.cass.pmc;NsClient++ CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1673745526] SERVICE ALERT: Veeam2019.cass.pmc;NsClient++ CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1673746786] SERVICE ALERT: Veeam2019.cass.pmc;NsClient++ CPU Load;OK;HARD;1;CPU Load 78% (5 min average)