[1716092239] SERVICE ALERT: SRV_nagios;Procesos;WARNING;SOFT;1;PROCS WARNING: 276 processes with STATE = RSZDT
[1716092299] SERVICE ALERT: SRV_nagios;Procesos;WARNING;SOFT;2;PROCS WARNING: 271 processes with STATE = RSZDT
[1716092359] SERVICE ALERT: SRV_nagios;Procesos;WARNING;SOFT;3;PROCS WARNING: 272 processes with STATE = RSZDT
[1716092419] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 271 processes with STATE = RSZDT
[1716107059] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.05, 0.78, 0.43
[1716107119] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.79, 0.65, 0.40
[1716108379] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.09, 1.11, 0.73
[1716108439] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.77, 1.32, 0.82
[1716108499] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.02, 1.08, 0.77
[1716109159] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.41, 1.28, 0.85
[1716109219] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.61, 1.19, 0.85
[1716120979] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.93, 0.65, 0.46
[1716121039] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.99, 0.64, 0.48
[1716127099] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.08, 0.69, 0.45
[1716127159] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.36, 0.79, 0.50
[1716133820] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.40, 1.29, 0.78
[1716133880] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.88, 1.06, 0.73
[1716146540] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.95, 1.07, 0.65
[1716146600] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.12, 0.89, 0.61
[1716149060] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.00, 0.98, 0.61
[1716149120] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.38, 0.91, 0.61
[1716150080] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.05, 0.73, 0.57
[1716150140] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.78, 0.61, 0.54
[1716151100] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.62, 0.81, 0.54
[1716151160] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.10, 0.72, 0.52
[1716154219] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 401 processes with STATE = RSZDT
[1716154519] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 393 processes with STATE = RSZDT
[1716156020] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 408 processes with STATE = RSZDT
[1716156320] SERVICE FLAPPING ALERT: SRV_nagios;Procesos;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1716156320] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 399 processes with STATE = RSZDT
[1716156920] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 402 processes with STATE = RSZDT
[1716157220] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 353 processes with STATE = RSZDT
[1716158720] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.48, 1.35, 0.86
[1716158780] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.91, 1.10, 0.81
[1716160220] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 407 processes with STATE = RSZDT
[1716160520] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 358 processes with STATE = RSZDT
[1716161240] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.17, 1.12, 0.69
[1716161300] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.88, 0.93, 0.65
[1716163520] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 408 processes with STATE = RSZDT
[1716163820] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 362 processes with STATE = RSZDT
[1716164120] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 405 processes with STATE = RSZDT
[1716164420] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 354 processes with STATE = RSZDT
[1716167120] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 419 processes with STATE = RSZDT
[1716167720] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 373 processes with STATE = RSZDT
[1716170720] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 422 processes with STATE = RSZDT
[1716171020] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 374 processes with STATE = RSZDT
[1716171320] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 423 processes with STATE = RSZDT
[1716171620] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 372 processes with STATE = RSZDT
[1716173060] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.81, 1.10, 0.73
[1716173120] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.64, 1.19, 0.78
[1716173720] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 432 processes with STATE = RSZDT
[1716174020] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 378 processes with STATE = RSZDT
[1716174320] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 430 processes with STATE = RSZDT
[1716174620] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 381 processes with STATE = RSZDT
[1716175220] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 439 processes with STATE = RSZDT
[1716175520] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 377 processes with STATE = RSZDT