[1683264539] HOST ALERT: Blueshell;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1683264540] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1683264540] SERVICE ALERT: FW Blueshell;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1683264540] SERVICE ALERT: Blueshell;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1683264570] HOST ALERT: Blueshell;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1683264629] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;2;(Host check timed out after 30.00 seconds)
[1683264660] HOST ALERT: Blueshell;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1683264719] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1683264750] HOST ALERT: Blueshell;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1683264809] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1683264840] HOST ALERT: Blueshell;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1683264899] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1683264931] HOST ALERT: Blueshell;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1683264989] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1683265021] HOST ALERT: Blueshell;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1683265079] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1683265111] HOST ALERT: Blueshell;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1683265169] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1683265201] HOST ALERT: Blueshell;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1683265259] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1683265291] HOST ALERT: Blueshell;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1683265349] HOST ALERT: FW Blueshell;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1683291321] HOST ALERT: SRV_sync.cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.2.26)
[1683291325] SERVICE ALERT: SRV_sync.cass.cl;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.2.26)
[1683291385] HOST ALERT: SRV_sync.cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1683291386] SERVICE ALERT: SRV_sync.cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1683294301] SERVICE ALERT: SRV_nagios;Actualizacion;OK;HARD;1;YUM OK: 0 Updates Available
[1683294371] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1683294492] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1683294614] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1683297244] SERVICE ALERT: SRV11VIBE_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1683297321] SERVICE ALERT: SRV11VIBE_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates CRITICAL : discontinued OS Release SLE 15.1: no updates available
[1683297570] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1683297693] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1683297816] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1683298193] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.
[1683306213] HOST ALERT: FW Blueshell;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 43.86 ms
[1683306214] SERVICE ALERT: Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 12.00 ms
[1683306214] SERVICE ALERT: FW Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 37.91 ms
[1683306215] HOST ALERT: Blueshell;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 8.20 ms
[1683306479] HOST ALERT: FW Blueshell;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1683306480] SERVICE ALERT: FW Blueshell;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1683306480] SERVICE ALERT: Blueshell;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1683306481] HOST ALERT: Blueshell;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1683306511] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1683306570] HOST ALERT: Blueshell;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1683306594] SERVICE ALERT: FW Blueshell;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 9.26 ms
[1683306594] HOST FLAPPING ALERT: FW Blueshell;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1683306594] HOST ALERT: FW Blueshell;UP;SOFT;1;PING OK - Packet loss = 79%, RTA = 8.21 ms
[1683306605] SERVICE ALERT: Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 7.58 ms
[1683306609] HOST ALERT: Blueshell;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 7.03 ms
[1683306634] SERVICE FLAPPING ALERT: FW Blueshell;Ping;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1683306634] SERVICE ALERT: FW Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 8.13 ms
[1683307114] HOST FLAPPING ALERT: FW Blueshell;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1683307774] SERVICE FLAPPING ALERT: FW Blueshell;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1683308130] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1683308250] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 11.61MB (9.29%), Average OUT: 224.63KB (0.18%)Total RX: 10478.93GBytes, Total TX: 6879.42GBytes
[1683312790] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1683312910] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1683317674] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1683317674] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1683317687] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1683317806] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1683317807] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1683317819] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1683318002] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1683318135] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1683318556] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1683318688] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available