[1713597002] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1713597122] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.071 sec. response time
[1713606730] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;1;YUM CRITICAL: 17 Updates Available
[1713606792] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;2;YUM CRITICAL: 17 Updates Available
[1713606855] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;3;YUM CRITICAL: 17 Updates Available
[1713606917] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;HARD;4;YUM CRITICAL: 17 Updates Available
[1713612002] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1713612122] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.021 sec. response time
[1713614462] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1713614471] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1713614533] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1713614582] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;OK;SOFT;2;c:\ - total: 99.33 Gb - used: 66.60 Gb (67%) - free 32.73 Gb (33%)
[1713622173] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 98% (5 min average)
[1713622293] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1713622413] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1713622533] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 94% (5 min average)
[1713634802] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1713634922] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.003 sec. response time
[1713639002] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1713639122] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.028 sec. response time
[1713640998] SERVICE ALERT: FW Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 37%, RTA = 5.98 ms
[1713640998] SERVICE ALERT: Blueshell;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 3.38 ms
[1713640998] HOST ALERT: FW Blueshell;UP;HARD;1;PING WARNING - Packet loss = 83%, RTA = 5.95 ms
[1713641008] HOST ALERT: Blueshell;UP;HARD;1;PING OK - Packet loss = 16%, RTA = 3.34 ms
[1713641039] SERVICE ALERT: Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.24 ms
[1713647402] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1713647522] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 43 day(s) 9 hour(s) 50 minute(s)
[1713659261] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;The command (check_updates) returned an invalid return code: 255
[1713659322] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1713664202] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1713664322] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.026 sec. response time
[1713669924] HOST ALERT: Clinica PV Telsur;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1713669932] SERVICE ALERT: ADL Diagnostic;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 52.10 ms
[1713669936] HOST ALERT: FW Clinica PV Movistar;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1713669941] SERVICE ALERT: FW Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 3.02 ms
[1713669942] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1713669942] SERVICE ALERT: SRV_email.cass.cl;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 0.90 ms
[1713669943] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 1.07 ms
[1713669944] SERVICE ALERT: FW Cass Telsur;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 90%, RTA = 0.53 ms
[1713669944] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1713669945] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1713669949] SERVICE ALERT: SRV11VIBE_cass.cl;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 0.88 ms
[1713669949] SERVICE ALERT: FW Clinica PV Movistar;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 82%, RTA = 44.06 ms
[1713669952] HOST ALERT: web_cass_cl;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1713669972] HOST ALERT: web_cass_cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1713669986] SERVICE ALERT: FW Clinica PV Movistar;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 34.36 ms
[1713669988] HOST ALERT: Clinica PV Telsur;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1713669991] HOST ALERT: FW Clinica PV Movistar;UP;SOFT;1;PING OK - Packet loss = 16%, RTA = 35.29 ms
[1713670056] SERVICE ALERT: ADL Diagnostic;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 38.55 ms
[1713670065] SERVICE ALERT: FW Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 24.55 ms
[1713670066] SERVICE ALERT: SRV_email.cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1713670067] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 18.89 ms
[1713670068] SERVICE ALERT: FW Cass Telsur;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1713670073] SERVICE ALERT: SRV11VIBE_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.93 ms
[1713670075] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1713670077] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1713670078] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1713670201] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1713670201] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1713670202] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1713670202] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1713670322] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 43 day(s) 16 hour(s) 10 minute(s)
[1713670334] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1713670334] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1713670334] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available