[1694495405] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694495537] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1694496069] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694496069] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694496129] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 143 day(s) 15 hour(s) 37 minute(s)
[1694496189] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.03, 0.02, 0.01
[1694507469] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694507469] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694507529] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 143 day(s) 18 hour(s) 47 minute(s)
[1694507589] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.01, 0.02, 0.01
[1694508669] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694508789] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.02, 0.01, 0.01
[1694512869] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694512929] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 143 day(s) 20 hour(s) 17 minute(s)
[1694514669] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694514789] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694515802] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694515869] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694515922] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.064 sec. response time
[1694515989] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.00, 0.01, 0.00
[1694516469] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694516589] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694516701] SERVICE ALERT: FW Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 81%, RTA = 5.40 ms
[1694516701] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1694516765] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 16.57 ms
[1694516825] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 5.26 ms
[1694518297] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518303] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518307] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518325] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518447] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518453] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518457] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518476] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518539] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 86%, RTA = 6.45 ms
[1694518597] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518603] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518607] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518626] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694518664] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 16%, RTA = 11.60 ms
[1694518802] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518803] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518804] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518804] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694518805] SERVICE ALERT: SRV_SOS2021.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694518809] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518812] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518816] SERVICE ALERT: SRV_Cfiler.cass.cl;Almacenamiento /;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518816] SERVICE ALERT: SRV_Cfiler.cass.cl;Https;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518818] SERVICE ALERT: SRV_SOS2021.cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694518818] SERVICE ALERT: SRV_email.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518818] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694518818] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518819] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518828] SERVICE ALERT: SRV_email.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518835] SERVICE ALERT: SRV_Cfiler.cass.cl;Procs Apache;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518838] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694518839] SERVICE ALERT: SRV_SOS2021.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694518850] SERVICE ALERT: SRV_Cfiler.cass.cl;Serv Apache;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518852] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518853] SERVICE ALERT: SRV_Cfiler.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518854] SERVICE ALERT: SRV_Cfiler.cass.cl;Http;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518857] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694518869] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694518917] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 146 day(s) 9 hour(s) 46 minute(s)
[1694518922] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.033 sec. response time
[1694518923] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 0 day(s) 19 hour(s) 29 minute(s)
[1694518924] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;OK;SOFT;2;POP OK - 0.001 second response time on 10.0.2.25 port 110 [+OK GroupWise POP3 server ready]
[1694518924] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694518925] SERVICE ALERT: SRV_SOS2021.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 0 users currently logged in
[1694518929] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;OK;SOFT;2;OK - load average: 0.00, 0.00, 0.02
[1694518932] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 3893 MB - Used: 986 MB - 25% [SWAP] Total: 2047 MB - Used: 429 MB - 20%
[1694518936] SERVICE ALERT: SRV_Cfiler.cass.cl;Almacenamiento /;OK;SOFT;2;DISK OK - free space: / 2465297 MB (73% inode=-):
[1694518936] SERVICE ALERT: SRV_Cfiler.cass.cl;Https;OK;SOFT;2;TCP OK - 0.000 second response time on port 443
[1694518938] SERVICE ALERT: SRV_SOS2021.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.01, 0.01, 0.01
[1694518938] SERVICE ALERT: SRV_email.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 0 users currently logged in
[1694518938] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 324 processes
[1694518938] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694518939] SERVICE ALERT: SRV_email.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 301 processes
[1694518949] SERVICE ALERT: SRV_email.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 11939 MB - Used: 1455 MB - 12% [SWAP] Total: 2048 MB - Used: 0 MB - 0%
[1694518956] SERVICE ALERT: SRV_Cfiler.cass.cl;Procs Apache;OK;SOFT;2;PROCS OK: 9 processes with args 'apache2'
[1694518958] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694518959] SERVICE ALERT: SRV_SOS2021.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7946 MB - Used: 1565 MB - 19% [SWAP] Total: 10239 MB - Used: 1 MB - 0%
[1694518971] SERVICE ALERT: SRV_Cfiler.cass.cl;Serv Apache;OK;SOFT;2;OK: Service apache is running!
[1694518972] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 308 processes
[1694518973] SERVICE ALERT: SRV_Cfiler.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 1 users currently logged in
[1694518974] SERVICE ALERT: SRV_Cfiler.cass.cl;Http;OK;SOFT;2;TCP OK - 0.000 second response time on port 80
[1694518989] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694519294] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 93%, RTA = 3.85 ms
[1694519402] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519403] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519404] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694519405] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519405] SERVICE ALERT: SRV_SOS2021.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694519409] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519412] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519416] SERVICE ALERT: SRV_Cfiler.cass.cl;Https;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519416] SERVICE ALERT: SRV_Cfiler.cass.cl;Almacenamiento /;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519418] SERVICE ALERT: SRV_email.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519418] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.25 ms
[1694519418] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694519418] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519418] SERVICE ALERT: SRV_SOS2021.cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694519419] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519429] SERVICE ALERT: SRV_email.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519436] SERVICE ALERT: SRV_Cfiler.cass.cl;Procs Apache;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519438] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694519439] SERVICE ALERT: SRV_SOS2021.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1694519451] SERVICE ALERT: SRV_Cfiler.cass.cl;Serv Apache;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519452] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519453] SERVICE ALERT: SRV_Cfiler.cass.cl;Usuarios Login;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519454] SERVICE ALERT: SRV_Cfiler.cass.cl;Http;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519457] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694519517] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1694519517] SERVICE ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 146 day(s) 9 hour(s) 56 minute(s)
[1694519522] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.006 sec. response time
[1694519523] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 0 day(s) 19 hour(s) 39 minute(s)
[1694519524] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694519525] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;OK;SOFT;2;POP OK - 0.000 second response time on 10.0.2.25 port 110 [+OK GroupWise POP3 server ready]
[1694519525] SERVICE ALERT: SRV_SOS2021.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 0 users currently logged in
[1694519529] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;OK;SOFT;2;OK - load average: 0.11, 0.05, 0.01
[1694519532] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 3893 MB - Used: 1003 MB - 25% [SWAP] Total: 2047 MB - Used: 429 MB - 20%
[1694519536] SERVICE ALERT: SRV_Cfiler.cass.cl;Https;OK;SOFT;2;TCP OK - 0.000 second response time on port 443
[1694519536] SERVICE ALERT: SRV_Cfiler.cass.cl;Almacenamiento /;OK;SOFT;2;DISK OK - free space: / 2465297 MB (73% inode=-):
[1694519538] SERVICE ALERT: SRV_email.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 0 users currently logged in
[1694519538] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 327 processes
[1694519538] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694519539] SERVICE ALERT: SRV_SOS2021.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.01, 0.02, 0.01
[1694519539] SERVICE ALERT: SRV_email.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 302 processes
[1694519549] SERVICE ALERT: SRV_email.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 11939 MB - Used: 1495 MB - 12% [SWAP] Total: 2048 MB - Used: 0 MB - 0%
[1694519556] SERVICE ALERT: SRV_Cfiler.cass.cl;Procs Apache;OK;SOFT;2;PROCS OK: 9 processes with args 'apache2'
[1694519558] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1694519560] SERVICE ALERT: SRV_SOS2021.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7946 MB - Used: 1561 MB - 19% [SWAP] Total: 10239 MB - Used: 1 MB - 0%
[1694519571] SERVICE ALERT: SRV_Cfiler.cass.cl;Serv Apache;OK;SOFT;2;OK: Service apache is running!
[1694519572] SERVICE ALERT: SRV_Cfiler.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 307 processes
[1694519573] SERVICE ALERT: SRV_Cfiler.cass.cl;Usuarios Login;OK;SOFT;2;USERS OK - 1 users currently logged in
[1694519574] SERVICE ALERT: SRV_Cfiler.cass.cl;Http;OK;SOFT;2;TCP OK - 0.000 second response time on port 80
[1694520002] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694520003] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694520003] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694520005] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1694520009] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694520009] SERVICE ALERT: SRV_chat.cass.cl;Procesos;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1694520012] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.20 port 5666: Connection refused
[1694520066] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1694520072] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1694520082] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1694520086] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1694520122] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.004 sec. response time
[1694520123] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 0 day(s) 19 hour(s) 49 minute(s)
[1694520125] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;OK;SOFT;2;POP OK - 0.000 second response time on 10.0.2.25 port 110 [+OK GroupWise POP3 server ready]
[1694520129] SERVICE ALERT: SRV_Cfiler.cass.cl;CPU;OK;SOFT;2;OK - load average: 0.21, 0.13, 0.04
[1694520129] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 310 processes
[1694520132] SERVICE ALERT: SRV_Cfiler.cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 3893 MB - Used: 978 MB - 25% [SWAP] Total: 2047 MB - Used: 429 MB - 20%
[1694520506] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694520626] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 29.57KB (0.02%), Average OUT: 13.26KB (0.01%)Total RX: 49.27GBytes, Total TX: 54.75GBytes
[1694520657] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Tiempo Activo;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1694521269] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.13. Check system logs on 10.0.2.13
[1694521389] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694522432] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694522553] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1694523463] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694523516] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694523583] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1694523636] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 24.19KB (0.02%), Average OUT: 55.62KB (0.04%)Total RX: 49.33GBytes, Total TX: 54.80GBytes
[1694524269] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694524269] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694524329] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 143 day(s) 23 hour(s) 27 minute(s)
[1694524389] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.03, 0.02, 0.01
[1694524546] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (210819931731.04%) too high
[1694524666] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;2;CRITICAL - IN bandwidth (314.98%) too high
[1694524786] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;3;OK - Average IN: 66.53KB (0.05%), Average OUT: 11.94KB (0.01%)Total RX: 49.37GBytes, Total TX: 54.82GBytes
[1694525469] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.13. Check system logs on 10.0.2.13
[1694525469] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694525529] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 143 day(s) 23 hour(s) 47 minute(s)
[1694525589] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.01, 0.02, 0.02
[1694529669] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694529729] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 0 hour(s) 57 minute(s)
[1694532427] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694532559] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1694533615] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1694533648] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694533735] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 49.20KB (0.04%), Average OUT: 522.69KB (0.42%)Total RX: 892.97GBytes, Total TX: 830.44GBytes
[1694533768] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 237.32B (0.00%), Average OUT: 89.13B (0.00%)Total RX: 11201.24GBytes, Total TX: 8360.15GBytes
[1694535069] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694535129] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 2 hour(s) 27 minute(s)
[1694535464] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1694535468] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;HARD;1;PING WARNING - Packet loss = 96%, RTA = 1.45 ms
[1694535502] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.07 ms
[1694535506] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1694535818] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694535836] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694535948] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1694535966] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1694535967] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694536068] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 314.12B (0.00%), Average OUT: 134.51B (0.00%)Total RX: 11201.24GBytes, Total TX: 8360.15GBytes
[1694536086] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;3;OK - Average IN: 10.97KB (0.01%), Average OUT: 9.29KB (0.01%)Total RX: 49.62GBytes, Total TX: 55.05GBytes
[1694536087] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1694536258] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694536394] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;(No output on stdout) stderr:
[1694536397] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694536397] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694536517] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 9.23KB (0.01%), Average OUT: 11.44KB (0.01%)Total RX: 49.62GBytes, Total TX: 55.05GBytes
[1694536517] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1694536565] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694536685] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 86.23KB (0.07%), Average OUT: 504.87KB (0.40%)Total RX: 893.14GBytes, Total TX: 831.65GBytes
[1694536793] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694536914] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1694536995] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694537115] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 60.21KB (0.05%), Average OUT: 338.97KB (0.27%)Total RX: 893.16GBytes, Total TX: 831.76GBytes
[1694537964] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694537965] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694538085] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 229.92B (0.00%), Average OUT: 80.74B (0.00%)Total RX: 11201.24GBytes, Total TX: 8360.15GBytes
[1694538085] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 26.00KB (0.02%), Average OUT: 162.26KB (0.13%)Total RX: 893.19GBytes, Total TX: 831.90GBytes
[1694539269] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694539389] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.12, 0.03, 0.02
[1694543114] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1694543124] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1694543507] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.52 Gb (80%) - free 23.88 Gb (20%)
[1694543627] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;SOFT;2;c:\ - total: 119.40 Gb - used: 95.51 Gb (80%) - free 23.89 Gb (20%)
[1694544047] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.87 Gb (80%) - free 23.53 Gb (20%)
[1694544069] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694544167] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.53 Gb (80%) - free 23.87 Gb (20%)
[1694544189] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694544287] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.53 Gb (80%) - free 23.86 Gb (20%)
[1694545067] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.50 Gb (80%) - free 23.90 Gb (20%)
[1694547406] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694547526] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1694548869] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694548869] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694548929] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 6 hour(s) 17 minute(s)
[1694548989] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.04, 0.03, 0.02
[1694549469] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694549590] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694549627] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.53 Gb (80%) - free 23.86 Gb (20%)
[1694549747] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.53 Gb (80%) - free 23.86 Gb (20%)
[1694549867] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.53 Gb (80%) - free 23.87 Gb (20%)
[1694550575] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694550664] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1694550687] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694550700] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 298.41B (0.00%), Average OUT: 103.70B (0.00%)Total RX: 11203.17GBytes, Total TX: 8369.88GBytes
[1694550784] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1694550807] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 9.54KB (0.01%), Average OUT: 14.42KB (0.01%)Total RX: 49.89GBytes, Total TX: 55.26GBytes
[1694551869] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694551989] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694552570] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694552696] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 295.82B (0.00%), Average OUT: 103.91B (0.00%)Total RX: 11203.17GBytes, Total TX: 8369.88GBytes
[1694553069] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694553069] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694553129] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 7 hour(s) 27 minute(s)
[1694553189] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.03, 0.01, 0.00
[1694553218] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1694553338] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 24.01KB (0.02%), Average OUT: 8.70KB (0.01%)Total RX: 49.92GBytes, Total TX: 55.28GBytes
[1694553602] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1694553722] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.012 sec. response time
[1694554248] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 95.42 Gb (80%) - free 23.98 Gb (20%)
[1694554269] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.13. Check system logs on 10.0.2.13
[1694554390] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.02, 0.02, 0.00
[1694555470] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694555590] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.02, 0.02, 0.00
[1694558469] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694558470] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694558530] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 8 hour(s) 57 minute(s)
[1694558590] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694561867] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;1;c:\ - total: 119.40 Gb - used: 95.63 Gb (80%) - free 23.76 Gb (20%)
[1694561987] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;SOFT;2;c:\ - total: 119.40 Gb - used: 95.63 Gb (80%) - free 23.77 Gb (20%)
[1694562107] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 95.63 Gb (80%) - free 23.77 Gb (20%)
[1694562670] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694562670] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694562730] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 10 hour(s) 7 minute(s)
[1694562790] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.03, 0.03, 0.02
[1694563870] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694563930] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 10 hour(s) 27 minute(s)
[1694565409] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1694565541] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1694565595] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694565657] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1694566202] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694566334] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1694571070] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694571190] SERVICE ALERT: SRV11VIBE_cass.cl;Proceso Vibe;OK;SOFT;2;PROCS OK: 0 processes with command name 'vibe'
[1694571279] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1694571340] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional
[1694572270] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694572270] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1694572330] SERVICE ALERT: SRV11VIBE_cass.cl;Tiempo Activo;WARNING;HARD;3;Uptime WARNING: 144 day(s) 12 hour(s) 47 minute(s)
[1694572390] SERVICE ALERT: SRV11VIBE_cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.04, 0.01, 0.01