[1684130840] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684130930] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684132040] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684132130] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684133000] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684133090] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1684133090] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684133540] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684133630] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684134320] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684134410] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684135580] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1684135820] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684135910] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684136180] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684136270] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1684136270] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684137412] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1684154248] HOST ALERT: FW Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1684154248] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1684154252] HOST ALERT: Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1684154254] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1684154282] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1684154342] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1684154372] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1684154432] HOST ALERT: Electric Works;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1684154462] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1684154522] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1684154552] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1684154612] HOST ALERT: Electric Works;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1684154642] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1684154702] HOST ALERT: Electric Works;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1684154732] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1684154792] HOST ALERT: Electric Works;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1684154822] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;8;(Host check timed out after 30.00 seconds)
[1684154882] HOST ALERT: Electric Works;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1684154912] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1684154960] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684154972] HOST ALERT: Electric Works;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1684155002] HOST ALERT: FW Electric Works;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1684155051] HOST ALERT: Electric Works;UP;SOFT;1;PING OK - Packet loss = 75%, RTA = 34.80 ms
[1684155068] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.15 ms
[1684155081] HOST ALERT: FW Electric Works;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1684155083] SERVICE ALERT: FW Electric Works;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 32.88 ms
[1684155110] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684155122] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 33.24 ms
[1684155126] HOST ALERT: FW Electric Works;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 39.79 ms
[1684155441] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684155531] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1684155531] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684155621] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684155711] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684155861] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684155944] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684155951] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684156075] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1684156201] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684156333] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1684156341] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684156431] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684156761] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684156851] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684157181] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684157271] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684157361] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684157451] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684157658] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684157662] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684157709] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684157713] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684157718] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684157803] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684157829] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 52.98B (0.00%), Average OUT: 572.77B (0.00%)Total RX: 14286.44GBytes, Total TX: 8884.31GBytes
[1684157838] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 58.16KB (0.05%), Average OUT: 450.59KB (0.36%)Total RX: 639.99GBytes, Total TX: 640.67GBytes
[1684157958] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684157959] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684157967] SERVICE ALERT: SRVCONT.cass.cl;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1684157968] SERVICE ALERT: SRVCONT.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1684157978] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1684157998] HOST ALERT: SRVCONT.cass.cl;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1684157998] SERVICE ALERT: SRVCONT.cass.cl;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1684157999] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;CRITICAL;HARD;1;CRITICAL - Socket timeout
[1684158002] SERVICE ALERT: SRV_chat.cass.cl;Puerto Portainer;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1684158006] SERVICE ALERT: SRVCONT.cass.cl;Almacenamiento;CRITICAL;HARD;1;connect to address 10.0.0.95 and port 12489: No route to host
[1684158028] HOST ALERT: SRVCONT.cass.cl;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.95)
[1684158088] SERVICE ALERT: SRVCONT.cass.cl;Memoria;OK;SOFT;2;Memory usage: total:15743.15 MB - used: 1705.00 MB (11%) - free: 14038.15 MB (89%)
[1684158088] SERVICE ALERT: SRVCONT.cass.cl;CPU;OK;SOFT;2;CPU Load 0% (5 min average)
[1684158090] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 30.44B (0.00%), Average OUT: 6.75B (0.00%)Total RX: 14286.44GBytes, Total TX: 8884.31GBytes
[1684158092] HOST ALERT: SRVCONT.cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 4.03 ms
[1684158092] SERVICE ALERT: SRVCONT.cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.79 ms
[1684158093] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;(No output on stdout) stderr:
[1684158098] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1684158109] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1684158122] SERVICE ALERT: SRV_chat.cass.cl;Puerto Portainer;OK;SOFT;2;TCP OK - 0.000 second response time on port 9443
[1684158123] SERVICE ALERT: SRVCONT.cass.cl;Almacenamiento;OK;HARD;1;c:\ - total: 299.33 Gb - used: 40.69 Gb (14%) - free 258.64 Gb (86%)
[1684158945] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1684159790] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1684159966] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;HARD;1;Import-Clixml : Fin inesperado de archivo al analizar Name. lnea 69063, posicin 53.
[1684160264] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 2 critical, 0 optional
[1684160386] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 2 critical, 0 optional
[1684160508] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1684161328] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;CRITICAL - IN bandwidth (210819924957.29%) too high
[1684161449] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;2;CRITICAL - IN bandwidth (4231.23%) too high
[1684161569] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;3;OK - Average IN: 189.08KB (0.15%), Average OUT: 156.19KB (0.12%)Total RX: 640.31GBytes, Total TX: 645.89GBytes
[1684162609] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1684162729] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1684162732] HOST ALERT: SRV_sync.cass.cl;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1684162735] SERVICE ALERT: SRV_sync.cass.cl;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1684162781] SERVICE ALERT: SRV_sync.cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 70%, RTA = 0.92 ms
[1684162785] HOST ALERT: SRV_sync.cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1684162870] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1684162899] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684162960] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684162963] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684162990] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1684163019] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 116.62KB (0.09%), Average OUT: 612.51KB (0.49%)Total RX: 640.49GBytes, Total TX: 646.39GBytes
[1684163086] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 68.59B (0.00%), Average OUT: 9.87B (0.00%)Total RX: 14286.44GBytes, Total TX: 8884.31GBytes
[1684163088] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 32.62KB (0.03%), Average OUT: 49.70KB (0.04%)Total RX: 47.61GBytes, Total TX: 25.54GBytes
[1684171789] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684171909] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 44.88KB (0.04%), Average OUT: 373.34KB (0.30%)Total RX: 640.93GBytes, Total TX: 654.18GBytes
[1684181819] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1684181939] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 46.65KB (0.04%), Average OUT: 367.83KB (0.29%)Total RX: 641.67GBytes, Total TX: 656.14GBytes
[1684188959] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1684189079] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1684190651] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684190771] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 2606 MB - 33% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1684191246] SERVICE ALERT: SRV05AD02_cass.cl;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 1.25 ms
[1684191370] SERVICE ALERT: SRV05AD02_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1684193297] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1684194251] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: No route to host
[1684194283] SERVICE ALERT: SRV01USER_cass.cl;Ping;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1684194286] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1684194293] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.12 port 5666: No route to host
[1684194347] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.58 ms
[1684194371] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 2607 MB - 33% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1684194380] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684194404] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 178.77 Gb (17%) - free 845.22 Gb (83%)
[1684194404] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;SOFT;2;Memory usage: total:23551.21 MB - used: 1649.82 MB (7%) - free: 21901.39 MB (93%)
[1684194404] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;SOFT;2;c:\ - total: 119.40 Gb - used: 89.06 Gb (75%) - free 30.34 Gb (25%)
[1684194404] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;SOFT;2;CPU Load 0% (5 min average)
[1684194404] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1684194407] SERVICE ALERT: SRV01USER_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1684194470] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1684194560] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1684194560] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684194623] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1684195451] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1684195571] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 2607 MB - 33% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1684195765] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1684202651] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1684202771] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 2617 MB - 34% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1684202867] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 96% (5 min average)
[1684202987] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1684203107] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1684203467] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 91% (5 min average)
[1684206251] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684206371] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 2620 MB - 34% [SWAP] Total: 7944 MB - Used: 0 MB - 0%
[1684206873] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1684206935] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1684209458] SERVICE ALERT: SRVCONT.cass.cl;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1684209458] SERVICE ALERT: SRVCONT.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1684209468] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;connect to address 10.0.0.95 and port 12489: Connection refused
[1684209468] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.95 port 5666: Connection refused
[1684209578] SERVICE ALERT: SRVCONT.cass.cl;Memoria;OK;SOFT;2;Memory usage: total:15743.15 MB - used: 1870.88 MB (12%) - free: 13872.27 MB (88%)
[1684209578] SERVICE ALERT: SRVCONT.cass.cl;CPU;OK;SOFT;2;CPU Load 7% (5 min average)
[1684209588] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)