[1716179494] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;1;[MEMORY] Total: 7740 MB - Used: 5430 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716179554] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;2;[MEMORY] Total: 7740 MB - Used: 5446 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716179614] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;3;[MEMORY] Total: 7740 MB - Used: 5429 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716179674] SERVICE ALERT: SRV_nagios;Memoria;WARNING;HARD;4;[MEMORY] Total: 7740 MB - Used: 5429 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716179780] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.74, 1.27, 0.78
[1716179840] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.05, 1.05, 0.74
[1716179974] SERVICE ALERT: SRV_nagios;Memoria;OK;HARD;1;[MEMORY] Total: 7740 MB - Used: 5415 MB - 69% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716180020] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 455 processes with STATE = RSZDT
[1716180274] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;1;[MEMORY] Total: 7740 MB - Used: 5434 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716180334] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;2;[MEMORY] Total: 7740 MB - Used: 5430 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716180394] SERVICE ALERT: SRV_nagios;Memoria;WARNING;SOFT;3;[MEMORY] Total: 7740 MB - Used: 5451 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716180454] SERVICE ALERT: SRV_nagios;Memoria;WARNING;HARD;4;[MEMORY] Total: 7740 MB - Used: 5468 MB - 70% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716180620] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 400 processes with STATE = RSZDT
[1716180920] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 442 processes with STATE = RSZDT
[1716181220] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 398 processes with STATE = RSZDT
[1716181520] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 451 processes with STATE = RSZDT
[1716182720] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 398 processes with STATE = RSZDT
[1716183020] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 403 processes with STATE = RSZDT
[1716183800] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.05, 1.23, 0.87
[1716183860] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.99, 1.11, 0.85
[1716184220] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 400 processes with STATE = RSZDT
[1716184220] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.16, 0.79, 0.72
[1716184280] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 3.22, 1.40, 0.94
[1716184340] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.43, 1.23, 0.92
[1716184520] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 402 processes with STATE = RSZDT
[1716190220] SERVICE FLAPPING ALERT: SRV_nagios;Procesos;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1716191120] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 387 processes with STATE = RSZDT
[1716192020] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 402 processes with STATE = RSZDT
[1716192320] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 396 processes with STATE = RSZDT
[1716192920] SERVICE FLAPPING ALERT: SRV_nagios;Procesos;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1716192920] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 402 processes with STATE = RSZDT
[1716193400] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.19, 0.93, 0.69
[1716193460] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.18, 0.89, 0.69
[1716194120] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 400 processes with STATE = RSZDT
[1716194420] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 401 processes with STATE = RSZDT
[1716195020] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 400 processes with STATE = RSZDT
[1716195320] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 402 processes with STATE = RSZDT
[1716196221] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 397 processes with STATE = RSZDT
[1716196821] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 401 processes with STATE = RSZDT
[1716197721] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 397 processes with STATE = RSZDT
[1716198021] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 419 processes with STATE = RSZDT
[1716198920] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.31, 0.91, 0.67
[1716198980] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.41, 0.90, 0.68
[1716199221] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 397 processes with STATE = RSZDT
[1716199821] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 407 processes with STATE = RSZDT
[1716199940] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.26, 0.89, 0.66
[1716200000] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.22, 1.21, 0.79
[1716200060] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.82, 0.99, 0.74
[1716202221] SERVICE ALERT: SRV_nagios;Procesos;WARNING;HARD;4;PROCS WARNING: 398 processes with STATE = RSZDT
[1716202355] SERVICE ALERT: SRV_nagios;Memoria;CRITICAL;HARD;4;[MEMORY] Total: 7740 MB - Used: 6206 MB - 80% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716202521] SERVICE ALERT: SRV_nagios;Procesos;CRITICAL;HARD;4;PROCS CRITICAL: 410 processes with STATE = RSZDT
[1716202655] SERVICE ALERT: SRV_nagios;Memoria;WARNING;HARD;4;[MEMORY] Total: 7740 MB - Used: 6193 MB - 79% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716202955] SERVICE ALERT: SRV_nagios;Memoria;CRITICAL;HARD;4;[MEMORY] Total: 7740 MB - Used: 6201 MB - 80% [SWAP] Total: 7935 MB - Used: 0 MB - 0%
[1716204020] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.88, 0.76, 0.57
[1716204080] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.84, 0.68, 0.55
[1716205040] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.86, 1.29, 0.72
[1716205100] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.05, 1.06, 0.68
[1716208221] SERVICE FLAPPING ALERT: SRV_nagios;Procesos;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1716209060] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.66, 1.07, 0.68
[1716209120] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.56, 1.06, 0.70
[1716210680] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.05, 1.09, 0.73
[1716210740] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.75, 0.89, 0.68
[1716214700] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.19, 0.97, 0.67
[1716214760] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.17, 0.91, 0.67
[1716215720] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 4.44, 1.39, 0.83
[1716215780] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.31, 1.41, 0.88
[1716215840] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.92, 1.17, 0.83
[1716218301] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.07, 0.82, 0.63
[1716218361] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.13, 0.81, 0.64
[1716222921] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.35, 0.95, 0.68
[1716222981] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.04, 0.85, 0.66
[1716225441] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.67, 0.77, 0.61
[1716225501] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.46, 1.08, 0.72
[1716225561] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.16, 0.97, 0.71
[1716226521] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.77, 1.65, 1.05
[1716226581] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.77, 1.49, 1.03
[1716227541] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.04, 0.80, 0.74
[1716227601] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.46, 0.89, 0.78
[1716228561] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.99, 1.19, 0.81
[1716228621] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.11, 1.17, 0.83
[1716228681] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.86, 0.99, 0.79
[1716229641] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.85, 1.02, 0.79
[1716229701] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.82, 0.89, 0.76
[1716230661] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.88, 1.21, 0.82
[1716230722] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.06, 0.99, 0.77
[1716233182] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.45, 1.50, 1.11
[1716233242] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.37, 1.39, 1.11
[1716234202] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.61, 1.42, 1.15
[1716234262] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;2;WARNING - load average: 2.00, 1.50, 1.19
[1716234322] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.73, 1.22, 1.12
[1716238282] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.68, 1.03, 0.71
[1716238342] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.99, 0.84, 0.67
[1716239302] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 4.79, 1.62, 0.91
[1716239362] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.02, 1.40, 0.88
[1716239422] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.82, 1.17, 0.83
[1716239783] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.48, 0.94, 0.76
[1716239843] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 3.11, 1.50, 0.97
[1716239903] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.34, 1.30, 0.93
[1716240863] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 4.74, 1.77, 1.05
[1716240923] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.23, 1.62, 1.05
[1716240983] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.91, 1.36, 0.99
[1716241343] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.89, 0.85, 0.83
[1716241403] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.94, 1.42, 1.03
[1716241463] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.16, 1.19, 0.97
[1716242423] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.51, 1.98, 1.36
[1716242483] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.97, 1.64, 1.28
[1716244943] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.07, 1.69, 1.27
[1716245003] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.39, 1.71, 1.30
[1716245063] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.11, 1.50, 1.25
[1716245423] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.73, 1.26, 1.12
[1716245483] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.52, 1.36, 1.16
[1716245543] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;3;WARNING - load average: 1.84, 1.50, 1.23
[1716245604] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;4;OK - load average: 0.67, 1.23, 1.15
[1716246564] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.67, 1.13, 0.90
[1716246624] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.01, 0.94, 0.85
[1716247584] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.55, 1.18, 0.83
[1716247644] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.00, 0.99, 0.79
[1716248604] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.32, 1.46, 0.93
[1716248664] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.85, 1.19, 0.87
[1716249024] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.93, 1.01, 0.83
[1716249084] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 2.76, 1.39, 0.97
[1716249144] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 1.30, 1.24, 0.94
[1716251604] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.35, 1.15, 0.84
[1716251664] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.86, 0.94, 0.79
[1716252624] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.14, 1.01, 0.78
[1716252684] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.83, 0.84, 0.74
[1716253644] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.85, 1.10, 0.80
[1716253704] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.15, 0.94, 0.77
[1716254664] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.32, 1.01, 0.81
[1716254724] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.57, 1.10, 0.85
[1716255684] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.19, 0.98, 0.75
[1716255744] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;2;WARNING - load average: 1.89, 1.17, 0.83
[1716255804] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;3;OK - load average: 0.69, 0.96, 0.78
[1716258264] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.00, 1.32, 0.87
[1716258324] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.10, 1.08, 0.81
[1716259284] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.38, 1.04, 0.77
[1716259344] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.87, 0.85, 0.73
[1716259704] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 4.51, 1.26, 0.82
[1716259764] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;2;CRITICAL - load average: 3.83, 1.75, 1.02
[1716259825] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;3;CRITICAL - load average: 2.21, 1.72, 1.06
[1716259885] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;4;OK - load average: 0.88, 1.43, 1.00
[1716260845] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 3.31, 1.48, 0.99
[1716260905] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 1.22, 1.21, 0.92
[1716261865] SERVICE ALERT: SRV_nagios;CPU;WARNING;SOFT;1;WARNING - load average: 1.96, 1.17, 0.90
[1716261925] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.72, 0.96, 0.84
[1716262885] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.65, 1.58, 0.99
[1716262945] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.97, 1.29, 0.93
[1716263905] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.42, 1.49, 1.06
[1716263965] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.94, 1.24, 1.00