[1682913885] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682914017] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682914202] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682914334] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682987193] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682987264] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682987326] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682987396] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682988002] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1682988076] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1682988134] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682988208] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1682992737] SERVICE ALERT: FW Cass Entel;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 82%, RTA = 38.99 ms
[1682992861] SERVICE ALERT: FW Cass Entel;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 46.15 ms
[1682993053] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1682993060] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1682993062] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1682993071] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1682993083] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1682993097] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1682993113] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1682993150] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1682993152] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1682993203] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1682993203] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1682993240] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1682993242] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1682993293] HOST ALERT: ADL Diagnostic;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1682993330] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1682993332] HOST ALERT: FW Cass Entel;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1682993383] HOST ALERT: ADL Diagnostic;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1682993420] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1682993422] HOST ALERT: FW Cass Entel;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1682993473] HOST ALERT: ADL Diagnostic;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1682993504] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1682993510] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1682993512] HOST ALERT: FW Cass Entel;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1682993563] HOST ALERT: ADL Diagnostic;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1682993600] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1682993602] HOST ALERT: FW Cass Entel;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1682993624] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1682993653] HOST ALERT: ADL Diagnostic;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1682993690] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1682993692] HOST ALERT: FW Cass Entel;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1682993743] HOST ALERT: ADL Diagnostic;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1682993744] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1682993780] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1682993782] HOST ALERT: FW Cass Entel;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1682993833] HOST ALERT: ADL Diagnostic;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1682993864] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 94% (5 min average)
[1682993870] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1682993872] HOST ALERT: FW Cass Entel;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1682997631] SERVICE ALERT: FW Cass Entel;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 40.67 ms
[1682997633] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 39.29 ms
[1682997633] HOST ALERT: ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 76%, RTA = 39.32 ms
[1682997633] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 76%, RTA = 39.25 ms
[1682997637] HOST ALERT: FW ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 40.51 ms
[1682997656] HOST ALERT: FW Cass Entel;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 37.77 ms
[1682997665] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 40.04 ms
[1682999714] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1682999836] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1682999957] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional