[1685340704] SERVICE ALERT: SRV_nagios;Actualizacion;WARNING;SOFT;1;YUM WARNING: Yum output signature does not match current known format. Please make sure you have upgraded to the latest version from https://github.com/HariSekhon/Nagios-Plugins. If the problem persists, please raise a ticket at https://github.com/HariSekhon/Nagios-Plugins/issues with the full -vvv output
[1685340767] SERVICE ALERT: SRV_nagios;Actualizacion;OK;SOFT;2;YUM OK: 0 Updates Available
[1685343695] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685343698] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685343698] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685343700] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685343827] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685343831] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685343831] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685343832] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685344201] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344201] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344201] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344333] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685344334] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685344334] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685344541] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344691] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;SOFT;2;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344841] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685344931] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685345021] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685345086] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1685345086] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.
[1685346222] SERVICE FLAPPING ALERT: SRVCONT.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1685347463] SERVICE ALERT: FW Confe;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 15642.79 ms
[1685347463] SERVICE ALERT: Confe;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 50%, RTA = 7486.77 ms
[1685347587] SERVICE ALERT: FW Confe;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 39.90 ms
[1685347587] SERVICE ALERT: Confe;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 39.51 ms
[1685373946] SERVICE ALERT: SRV_email.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 371 processes
[1685374003] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685374066] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;SOFT;2;PROCS CRITICAL: 387 processes
[1685374113] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;CRITICAL;SOFT;1;connect to address 10.0.2.25 and port 110: Connection refused
[1685374128] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685374142] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;connect to address 10.0.2.25 and port 25: Connection refused
[1685374186] SERVICE ALERT: SRV_email.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 371 processes
[1685374233] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;CRITICAL;SOFT;2;connect to address 10.0.2.25 and port 110: Connection refused
[1685374262] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;2;connect to address 10.0.2.25 and port 25: Connection refused
[1685374353] SERVICE ALERT: SRV_email.cass.cl;Correo Pop;OK;SOFT;3;POP OK - 0.002 second response time on 10.0.2.25 port 110 [+OK GroupWise POP3 server ready]
[1685374366] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 398 processes
[1685374382] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;3;SMTP OK - 0.206 sec. response time
[1685374825] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1685374945] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1685375326] SERVICE ALERT: SRV_email.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 379 processes
[1685375446] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 383 processes
[1685375746] SERVICE ALERT: SRV_email.cass.cl;Procesos;OK;HARD;1;PROCS OK: 348 processes
[1685375927] SERVICE ALERT: Servidores CASS;Estado Puerto 80;CRITICAL;SOFT;1;connect to address 10.0.0.17 and port 80: No route to host
[1685375946] SERVICE ALERT: Servidores CASS;Estado Puerto 3;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "10.0.0.17"
[1685375948] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1685375964] SERVICE ALERT: Servidores CASS;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1685376047] SERVICE ALERT: Servidores CASS;Estado Puerto 80;CRITICAL;SOFT;2;connect to address 10.0.0.17 and port 80: No route to host
[1685376077] SERVICE ALERT: Servidores CASS;Estado Puerto 3;WARNING;SOFT;2;WARNING: SNMP error: No response from remote host "10.0.0.17"
[1685376078] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1685376095] SERVICE ALERT: Servidores CASS;Tiempo Activo;CRITICAL;SOFT;2;CRITICAL - Plugin timed out while executing system call
[1685376167] SERVICE ALERT: Servidores CASS;Estado Puerto 80;CRITICAL;HARD;3;connect to address 10.0.0.17 and port 80: No route to host
[1685376207] SERVICE ALERT: Servidores CASS;Estado Puerto 3;WARNING;HARD;3;WARNING: SNMP error: No response from remote host "10.0.0.17"
[1685376209] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;(No output on stdout) stderr:
[1685376226] SERVICE ALERT: Servidores CASS;Tiempo Activo;CRITICAL;HARD;3;CRITICAL - Plugin timed out while executing system call
[1685376706] SERVICE ALERT: SRV_email.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 359 processes
[1685376826] SERVICE ALERT: SRV_email.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 346 processes
[1685377007] SERVICE ALERT: Servidores CASS;Tiempo Activo;OK;HARD;1;SNMP OK - Timeticks: (587587069) 68 days, 0:11:10.69
[1685377007] SERVICE ALERT: Servidores CASS;Estado Puerto 80;OK;HARD;1;TCP OK - 0.002 second response time on 10.0.0.17 port 80
[1685377048] SERVICE ALERT: Servidores CASS;Estado Puerto 3;OK;HARD;1;OK: Interface GigabitEthernet1/0/3 (index 3) is up.
[1685377049] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 56.92KB (0.05%), Average OUT: 75.69KB (0.06%)Total RX: 17906.47GBytes, Total TX: 48959.10GBytes
[1685387995] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1685388116] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 31.84B (0.00%), Average OUT: 5.27B (0.00%)Total RX: 726.54GBytes, Total TX: 1279.93GBytes
[1685388215] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1685388335] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1685398802] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1685398922] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.005 sec. response time
[1685404802] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1685404922] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1685406001] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685406121] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;OK;SOFT;2;PROCS OK: 3 processes with command name 'nginx'
[1685407202] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1685407202] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1685407322] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1817879 MiB (92.83% inode=99%):
[1685407322] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1685407397] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685407398] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685407529] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685407531] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685411402] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685411402] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1685411522] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1817865 MiB (92.83% inode=99%):
[1685411522] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1685412749] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 99% (5 min average)
[1685412869] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1685412989] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1685413169] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 90% (5 min average)
[1685413202] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1685413322] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1817862 MiB (92.83% inode=99%):
[1685414401] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685414521] SERVICE ALERT: SRV_chat.cass.cl;Proceso Nginx;OK;SOFT;2;PROCS OK: 3 processes with command name 'nginx'
[1685415602] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685415722] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1685416451] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685416512] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional