[1679555814] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679555955] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679556635] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679556776] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679556901] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557005] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557041] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557146] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557191] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557341] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557346] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557356] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557490] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557491] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557506] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557573] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679557641] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557656] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557700] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557714] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557775] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557795] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679557845] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557864] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679557926] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679557995] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558074] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558076] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558105] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558141] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679558220] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679558224] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558249] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679558351] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558366] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679558459] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558490] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679558516] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558550] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558608] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679558666] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558699] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679558758] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558811] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679558849] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558880] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558908] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679558998] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559030] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559054] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679559081] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559148] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559180] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559231] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559263] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679559264] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559298] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559330] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559379] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679559408] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559448] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559471] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559558] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559591] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1679559591] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679559649] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559658] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559681] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559683] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559798] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559799] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559820] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679559943] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679559948] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679559970] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560093] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560098] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560120] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560243] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560247] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679560266] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679560373] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560391] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679560517] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560520] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679560596] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560601] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560667] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560670] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560742] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679560749] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679560811] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679560817] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679560892] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560899] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560961] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679560967] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561031] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679561049] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561111] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561116] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679561197] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679561260] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679561266] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561347] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561361] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561410] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561416] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561497] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561510] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679561557] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679561562] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679561642] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679561707] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561712] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561720] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561792] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561849] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679561862] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561870] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679561942] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562012] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562013] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679562059] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562082] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679562163] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562203] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679562232] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562304] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679562353] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562382] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562454] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562501] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679562532] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562602] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679562618] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679562637] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679562651] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562682] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562752] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562762] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562801] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562832] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562897] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679562912] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562951] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679562982] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563047] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563062] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563071] SERVICE FLAPPING ALERT: SRV_email.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1679563071] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679563197] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563329] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1679563329] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679563347] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563512] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563581] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1679563662] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563807] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679563852] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679563957] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564003] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564107] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564153] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564257] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564646] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679564707] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679564783] SERVICE FLAPPING ALERT: SRV_email.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1679564857] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679565007] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679565177] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 381 processes
[1679565298] SERVICE ALERT: SRV_email.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 380 processes
[1679565454] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679565517] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679565538] SERVICE ALERT: SRV_email.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 381 processes
[1679565667] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679565817] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679566100] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679566161] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679566312] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679566462] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679566977] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679567217] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679567277] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679567427] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679567563] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1679567577] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679567647] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679567797] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679567947] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679568756] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679568817] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679568967] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569117] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569155] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679569219] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679569223] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569287] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569367] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679569431] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679569437] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569492] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569520] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569587] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569642] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569670] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569792] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569820] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679569867] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679569987] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570137] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570287] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1679570287] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570323] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679570359] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679570390] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570422] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570540] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570572] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570653] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679570690] SERVICE FLAPPING ALERT: SRV_email.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1679570690] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679570722] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1679570722] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571277] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571427] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571549] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679571577] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571620] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571716] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679571770] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571777] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571885] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679571901] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679571920] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571927] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571952] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679571967] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572077] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572102] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572115] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679572252] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572265] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572415] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572565] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572640] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679572702] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572852] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679572858] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1679572978] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 350 processes
[1679573002] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573090] SERVICE FLAPPING ALERT: SRV_email.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1679573098] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1679573218] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 351 processes
[1679573242] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679573307] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573329] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679573338] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 351 processes
[1679573390] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573398] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 344 processes
[1679573445] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679573457] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573512] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573541] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573607] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573662] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573675] SERVICE FLAPPING ALERT: SRV_Cfiler.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1679573691] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573812] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679573818] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1679573938] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 351 processes
[1679573971] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1679574058] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 353 processes
[1679574091] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1679574311] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679574935] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679574982] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1679575085] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679575235] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679575703] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 5.54 ms
[1679575827] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 11.37 ms
[1679577357] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1679577477] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 353 processes
[1679577597] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 354 processes
[1679577621] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679577682] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679577832] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679577959] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679577982] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578243] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679578258] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 349 processes
[1679578271] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679578304] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679578461] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578611] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679578617] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1679578677] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578708] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1679578732] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578737] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 348 processes
[1679578827] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578877] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679578977] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679578978] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1679579001] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579062] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679579098] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 356 processes
[1679579147] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579152] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579166] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579218] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 355 processes
[1679579286] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679579290] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679579316] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579427] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579436] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579440] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579458] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 347 processes
[1679579466] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679579573] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679579580] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679579590] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579699] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1679579735] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679579790] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579819] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 357 processes
[1679579843] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579885] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679579937] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679579939] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;3;PROCS OK: 348 processes
[1679579986] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679580035] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580119] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 356 processes
[1679580156] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580176] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679580207] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580239] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 356 processes
[1679580302] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679580357] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580359] SERVICE FLAPPING ALERT: SRV_chat.cass.cl;Procesos;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1679580359] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 355 processes
[1679580446] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580496] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580507] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580512] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580539] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 349 processes
[1679580584] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679580595] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679580642] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679580649] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679580779] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 358 processes
[1679580837] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580860] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580865] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580899] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 354 processes
[1679580912] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679580982] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581010] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581015] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581019] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;3;PROCS OK: 349 processes
[1679581053] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581158] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679581163] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679581192] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581199] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1679581308] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581319] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1679581319] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 351 processes
[1679581342] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581379] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 391 processes
[1679581439] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 355 processes
[1679581443] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581447] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581499] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 349 processes
[1679581591] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581612] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581613] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581739] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 351 processes
[1679581754] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581763] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679581801] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581859] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 349 processes
[1679581913] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581951] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679581964] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582063] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582099] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 356 processes
[1679582101] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582115] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582213] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582219] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 349 processes
[1679582249] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679582257] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582265] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582299] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679582311] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582405] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679582460] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1679582461] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582471] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1679582471] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679582543] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582580] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 348 processes
[1679582610] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679582655] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582675] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582693] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582760] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1679582760] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582806] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582825] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679582828] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679582880] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 348 processes
[1679582909] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679582953] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679582969] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679582978] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583119] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583126] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679583179] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583180] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1679583223] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583269] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583300] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 354 processes
[1679583329] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583336] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583373] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583419] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583420] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;3;PROCS OK: 348 processes
[1679583479] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583486] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583523] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583541] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 353 processes
[1679583612] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679583636] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583661] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 356 processes
[1679583673] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583781] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 355 processes
[1679583803] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679583810] SERVICE FLAPPING ALERT: SRV_chat.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1679583810] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679583823] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583835] SERVICE FLAPPING ALERT: SRV_email.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1679583835] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679583870] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583901] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;HARD;1;PROCS OK: 349 processes
[1679583966] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583973] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679583989] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584020] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584061] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1679584111] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584139] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584141] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1679584243] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584261] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 354 processes
[1679584287] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679584290] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584321] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584382] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;3;PROCS OK: 348 processes
[1679584385] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584433] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584437] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584471] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584535] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584562] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1679584583] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584587] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584614] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;3;Updates OK : no updates available
[1679584675] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1679584681] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1679584682] SERVICE ALERT: SRV_chat.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 349 processes
[1679584733] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;2;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584795] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1679584797] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584824] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584831] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1679584862] SERVICE ALERT: SRV_chat.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1679584883] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681411173] SERVICE FLAPPING ALERT: SRV_chat.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1681407608] HOST ALERT: Blueshell;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1681407619] SERVICE ALERT: SRV_email.cass.cl;Procesos;OK;HARD;1;PROCS OK: 295 processes
[1681407621] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;1;(Host check timed out after 30.02 seconds)
[1681407625] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STARTED; Service appears to have started flapping (30.0% change >= 20.0% threshold)
[1681407625] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681407634] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681407636] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1681407636] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681407638] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;WARNING;SOFT;1;c:\ - total: 99.33 Gb - used: 79.95 Gb (80%) - free 19.38 Gb (20%)
[1681407641] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681407651] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;HARD;1;Updates OK : no updates available
[1681407654] SERVICE ALERT: Blueshell;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1681407655] SERVICE ALERT: FW Blueshell;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1681407662] SERVICE ALERT: SRV_chat.cass.cl;Procesos;CRITICAL;SOFT;2;PROCS CRITICAL: 439 processes
[1681407662] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1681407666] SERVICE ALERT: SRV_nagios;CPU;CRITICAL;SOFT;1;CRITICAL - load average: 2.45, 0.77, 0.27
[1681407698] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1681407723] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 3 critical, 0 optional
[1681407726] SERVICE ALERT: SRV_nagios;CPU;OK;SOFT;2;OK - load average: 0.97, 0.66, 0.26
[1681407758] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;WARNING;SOFT;2;c:\ - total: 99.33 Gb - used: 79.95 Gb (80%) - free 19.38 Gb (20%)
[1681407782] SERVICE ALERT: SRV_chat.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 439 processes
[1681407789] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;3;(Host check timed out after 30.59 seconds)
[1681407878] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;WARNING;HARD;3;c:\ - total: 99.33 Gb - used: 79.95 Gb (80%) - free 19.38 Gb (20%)
[1681407879] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;4;PING CRITICAL - Packet loss = 100%
[1681407969] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1681408059] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1681408149] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1681408239] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1681408329] HOST ALERT: FW Blueshell;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1681408419] HOST ALERT: FW Blueshell;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1681408537] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1681408790] SERVICE FLAPPING ALERT: SRV_chat.cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1681409199] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681409215] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681409224] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681409228] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681409231] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681409259] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681409283] HOST ALERT: FW Blueshell;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 4.68 ms
[1681409308] SERVICE ALERT: Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.19 ms
[1681409309] SERVICE ALERT: FW Blueshell;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.21 ms
[1681409347] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681409356] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681409359] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681409364] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681411212] HOST ALERT: Blueshell;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 3.16 ms
[1681413388] SERVICE ALERT: SRV_nagios;Actualizacion;OK;HARD;1;YUM OK: 0 Updates Available
[1681413399] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681413459] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 388 processes
[1681413519] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681413579] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Procesos;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1681413579] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681413639] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681413699] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681413759] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681413819] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 387 processes
[1681414119] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681414179] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681414239] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681414299] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 387 processes
[1681414360] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681414420] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681414480] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681414537] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681414540] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681414546] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681414549] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681414600] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681414660] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 398 processes
[1681414668] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681414678] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681414681] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681414720] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681414781] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 393 processes
[1681414841] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681414901] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 392 processes
[1681414932] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1681414961] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681415021] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681415052] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1681415202] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681415262] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681415322] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681415382] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681415442] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681415502] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681415683] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681415743] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 391 processes
[1681415803] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681415863] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 391 processes
[1681415923] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681415983] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681416163] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681416223] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 394 processes
[1681417003] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681417064] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681417124] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 410 processes
[1681417184] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681417244] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 407 processes
[1681417304] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681417364] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681417424] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681417484] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681417545] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681417605] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681417665] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681417725] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681417785] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 390 processes
[1681417845] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681417905] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 393 processes
[1681417966] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 408 processes
[1681418026] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 393 processes
[1681418566] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 414 processes
[1681418626] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 394 processes
[1681418686] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 411 processes
[1681418746] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 395 processes
[1681418806] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 407 processes
[1681418867] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 389 processes
[1681418927] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681418987] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 394 processes
[1681419047] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681419107] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 391 processes
[1681419167] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 408 processes
[1681419227] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 391 processes
[1681419287] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681419348] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 393 processes
[1681419648] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 406 processes
[1681419708] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681419948] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681420008] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681420249] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681420309] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 395 processes
[1681420369] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681420429] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681420849] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681420909] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681420969] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681421029] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681421089] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681421150] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681421210] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681421270] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681421330] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681421390] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681421450] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681421630] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681421690] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681421750] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681421810] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681421870] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681421930] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681421990] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681422051] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 405 processes
[1681422111] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681422171] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 404 processes
[1681422231] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 395 processes
[1681422291] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681422351] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681422411] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 403 processes
[1681422471] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681422532] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681422592] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 395 processes
[1681422652] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 402 processes
[1681422712] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681422772] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 411 processes
[1681422832] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 399 processes
[1681422892] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 412 processes
[1681422953] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 398 processes
[1681423013] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 412 processes
[1681423073] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 396 processes
[1681423133] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681423193] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 397 processes
[1681423253] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681423313] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 398 processes
[1681423553] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;CRITICAL;HARD;3;PROCS CRITICAL: 401 processes
[1681423613] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 399 processes
[1681424754] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Procesos;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1681431678] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681431688] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681431691] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681431714] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681431809] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681431820] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681431823] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681431846] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681442050] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681442053] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681442076] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681442099] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1681442182] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681442185] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681442208] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1681442231] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available