[1705212043] HOST ALERT: Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1705212073] HOST ALERT: FW Gore Los Lagos;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1705212077] SERVICE ALERT: Gore Los Lagos;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1705212086] HOST ALERT: Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 64%, RTA = 6.20 ms
[1705212090] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 6.42 ms
[1705212111] SERVICE ALERT: Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 3.68 ms
[1705230292] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1705230293] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1705230293] SERVICE ALERT: SRV01USER_cass.cl;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.0.12)
[1705230301] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1705230306] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1705230306] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1705230307] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;HARD;3;connect to address 10.0.0.12 and port 12489: Connection refused
[1705230350] SERVICE ALERT: SRV01USER_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1705230354] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1705230362] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1705230366] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;HARD;1;CPU Load 0% (5 min average)
[1705230366] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;HARD;1;Memory usage: total:23551.21 MB - used: 1809.36 MB (8%) - free: 21741.85 MB (92%)
[1705230367] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 97.52 Gb (82%) - free 21.88 Gb (18%)
[1705230412] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 239.60 Gb (23%) - free 784.38 Gb (77%)
[1705237776] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;1;YUM CRITICAL: Yum nagios plugin has self terminated after exceeding the timeout (30 seconds)
[1705237838] SERVICE ALERT: SRV_nagios;Actualizacion;OK;SOFT;2;YUM OK: 0 Updates Available
[1705242840] SERVICE ALERT: SRV11VIBE_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1705242917] SERVICE ALERT: SRV11VIBE_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates CRITICAL : discontinued OS Release SLE 15.1: At least one of your Repositories might be out of date. Please run "zypper refresh" as root to update it.  3 security update(s): 1 recommended update(s):  2 package update(s):
[1705257584] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1705257717] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1705258201] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1705258334] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1705280652] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.92%) too high
[1705280773] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 46.53MB (37.22%), Average OUT: 75.98MB (60.78%)Total RX: 11946.50GBytes, Total TX: 49457.99GBytes
[1705285024] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1705285145] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1705286148] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1705286212] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 4 critical, 0 optional