[1674875184] SERVICE ALERT: Blueshell;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 3.22 ms
[1674875308] SERVICE ALERT: Blueshell;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.84 ms
[1674882251] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.28, 0.09, 0.04
[1674882371] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.26, 0.17, 0.07
[1674882491] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.13, 0.15, 0.08
[1674882551] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;HARD;1;OK - load average per CPU: 0.06, 0.13, 0.08
[1674886631] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.16, 0.15, 0.10
[1674886751] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1674886871] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.15, 0.14, 0.11
[1674893291] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;The command (check_updates) returned an invalid return code: 255
[1674893415] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1674903011] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.43, 0.24, 0.18
[1674903551] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.02, 0.17, 0.20
[1674904331] SERVICE ALERT: SRV_email.cass.cl;Carga CPU;OK;HARD;1;OK - load average per CPU: 0.04, 0.04, 0.10
[1674906045] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1674906195] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;2;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1674906345] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1674906408] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 1 optional
[1674913564] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (93.54%) too high
[1674913684] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 17.23MB (13.78%), Average OUT: 104.04MB (83.23%)Total RX: 8636.69GBytes, Total TX: 16371.46GBytes
[1674913804] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (94.45%) too high
[1674913925] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.49%) too high
[1674914045] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (91.39%) too high
[1674914225] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 23.03MB (18.42%), Average OUT: 97.21MB (77.77%)Total RX: 8640.30GBytes, Total TX: 16432.48GBytes
[1674914826] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (91.33%) too high
[1674914946] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (95.21%) too high
[1674915066] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (88.72%) too high
[1674915127] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1674915127] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 30.47MB (24.38%), Average OUT: 96.14MB (76.91%)Total RX: 8657.87GBytes, Total TX: 16524.50GBytes
[1674915307] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (91.09%) too high
[1674915427] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (95.76%) too high
[1674915548] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.36%) too high
[1674915788] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;CRITICAL;HARD;3;CRITICAL - OUT bandwidth (98.58%) too high
[1674915848] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 20.76MB (16.61%), Average OUT: 101.89MB (81.51%)Total RX: 8664.31GBytes, Total TX: 16605.52GBytes
[1674916208] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (95.55%) too high
[1674916329] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (93.23%) too high
[1674916449] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 10.33MB (8.27%), Average OUT: 100.28MB (80.22%)Total RX: 8674.35GBytes, Total TX: 16667.59GBytes
[1674916749] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.57%) too high
[1674916870] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 15.49MB (12.39%), Average OUT: 100.08MB (80.06%)Total RX: 8680.17GBytes, Total TX: 16709.93GBytes
[1674917230] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (88.72%) too high
[1674917351] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 10.21MB (8.17%), Average OUT: 103.86MB (83.08%)Total RX: 8689.50GBytes, Total TX: 16757.63GBytes
[1674917531] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.35%) too high
[1674917652] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 16.95MB (13.56%), Average OUT: 100.04MB (80.03%)Total RX: 8693.60GBytes, Total TX: 16788.47GBytes
[1674917832] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1674918493] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.54%) too high
[1674918613] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.45%) too high
[1674918733] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;3;OK - Average IN: 7.32MB (5.86%), Average OUT: 103.70MB (82.96%)Total RX: 8721.80GBytes, Total TX: 16884.49GBytes
[1674918913] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.45%) too high
[1674919033] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 11.94MB (9.55%), Average OUT: 98.47MB (78.78%)Total RX: 8725.33GBytes, Total TX: 16914.36GBytes
[1674919994] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.25%) too high
[1674920114] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 28.29MB (22.63%), Average OUT: 87.10MB (69.68%)Total RX: 8747.02GBytes, Total TX: 17019.28GBytes
[1674928874] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.22%) too high
[1674928994] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 13.13MB (10.50%), Average OUT: 105.81MB (84.65%)Total RX: 8969.05GBytes, Total TX: 17824.28GBytes
[1674929354] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.25%) too high
[1674929474] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 12.59MB (10.08%), Average OUT: 98.79MB (79.03%)Total RX: 8975.31GBytes, Total TX: 17872.97GBytes
[1674929714] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.07%) too high
[1674929835] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 13.11MB (10.49%), Average OUT: 102.83MB (82.27%)Total RX: 8979.91GBytes, Total TX: 17909.45GBytes
[1674930135] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.57%) too high
[1674930256] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 13.04MB (10.44%), Average OUT: 103.74MB (82.99%)Total RX: 8985.44GBytes, Total TX: 17953.21GBytes
[1674930436] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.43%) too high
[1674930557] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 12.17MB (9.74%), Average OUT: 102.48MB (81.99%)Total RX: 8989.24GBytes, Total TX: 17984.57GBytes
[1674941442] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1674941574] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1674954317] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.30%) too high
[1674954438] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 25.06MB (20.05%), Average OUT: 85.65MB (68.52%)Total RX: 9044.32GBytes, Total TX: 18593.79GBytes
[1674954738] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.95%) too high
[1674954858] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.37%) too high
[1674954979] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (87.47%) too high
[1674955039] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 1.18MB (0.95%), Average OUT: 99.10MB (79.28%)Total RX: 9049.09GBytes, Total TX: 18652.09GBytes
[1674955219] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.12%) too high
[1674955339] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 1.47MB (1.18%), Average OUT: 102.76MB (82.21%)Total RX: 9049.53GBytes, Total TX: 18682.85GBytes
[1674955699] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.39%) too high
[1674955819] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.57%) too high
[1674955940] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (88.57%) too high
[1674956241] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1674956241] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 4.28MB (3.42%), Average OUT: 99.16MB (79.33%)Total RX: 9051.06GBytes, Total TX: 18780.16GBytes
[1674956300] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (92.73%) too high
[1674956420] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.86%) too high
[1674956541] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (90.74%) too high
[1674956601] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 3.33MB (2.66%), Average OUT: 105.87MB (84.70%)Total RX: 9052.48GBytes, Total TX: 18820.26GBytes
[1674956661] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (90.40%) too high
[1674956782] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 4.02MB (3.22%), Average OUT: 106.22MB (84.98%)Total RX: 9053.18GBytes, Total TX: 18839.89GBytes
[1674956902] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.32%) too high
[1674957022] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.69%) too high
[1674957143] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (87.24%) too high
[1674957324] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 1.26MB (1.01%), Average OUT: 101.89MB (81.51%)Total RX: 9055.09GBytes, Total TX: 18899.03GBytes
[1674957383] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (94.58%) too high
[1674957503] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (87.24%) too high
[1674957624] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (88.41%) too high
[1674957684] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 7.62MB (6.10%), Average OUT: 104.59MB (83.67%)Total RX: 9056.79GBytes, Total TX: 18938.74GBytes
[1674957864] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.49%) too high
[1674957984] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 3.64MB (2.92%), Average OUT: 101.75MB (81.40%)Total RX: 9057.65GBytes, Total TX: 18970.15GBytes
[1674958164] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.87%) too high
[1674958284] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 4.72MB (3.78%), Average OUT: 105.98MB (84.79%)Total RX: 9058.89GBytes, Total TX: 19002.14GBytes
[1674958704] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.63%) too high
[1674958825] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 11.72MB (9.37%), Average OUT: 92.62MB (74.09%)Total RX: 9066.18GBytes, Total TX: 19052.21GBytes
[1674959305] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.19%) too high
[1674959426] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 15.19MB (12.15%), Average OUT: 102.64MB (82.11%)Total RX: 9072.73GBytes, Total TX: 19109.31GBytes
[1674959546] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1674959906] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.48%) too high
[1674960026] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 566.40KB (0.45%), Average OUT: 104.15MB (83.32%)Total RX: 9079.55GBytes, Total TX: 19168.24GBytes
[1674960147] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (87.60%) too high
[1674960267] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (94.53%) too high
[1674960387] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (90.59%) too high
[1674960507] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 870.30KB (0.70%), Average OUT: 103.23MB (82.58%)Total RX: 9079.87GBytes, Total TX: 19221.86GBytes
[1674960567] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (94.63%) too high
[1674960688] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;SOFT;2;OK - Average IN: 551.22KB (0.44%), Average OUT: 102.69MB (82.15%)Total RX: 9079.97GBytes, Total TX: 19241.39GBytes
[1674960988] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (86.44%) too high
[1674961108] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (88.62%) too high