[1675656367] HOST ALERT: FW ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1675656370] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1675656371] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1675656373] HOST ALERT: FW Confe;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1675656375] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1675656386] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1675656406] SERVICE ALERT: Confe;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1675656410] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1675656412] SERVICE ALERT: FW Confe;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1675656417] HOST ALERT: Confe;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1675656440] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1675656447] HOST ALERT: FW Confe;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1675656460] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1675656476] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1675656496] HOST ALERT: Confe;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1675656530] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1675656536] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1675656537] HOST ALERT: FW Confe;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1675656550] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1675656556] SERVICE ALERT: Confe;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1675656566] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1675656586] HOST ALERT: Confe;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1675656620] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1675656627] HOST ALERT: FW Confe;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1675656640] HOST ALERT: FW Cass Entel;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1675656656] HOST ALERT: ADL Diagnostic;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1675656676] HOST ALERT: Confe;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1675656710] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1675656717] HOST ALERT: FW Confe;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1675656730] HOST ALERT: FW Cass Entel;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1675656746] HOST ALERT: ADL Diagnostic;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1675656766] HOST ALERT: Confe;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1675656800] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1675656807] HOST ALERT: FW Confe;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1675656820] HOST ALERT: FW Cass Entel;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1675656836] HOST ALERT: ADL Diagnostic;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1675656856] HOST ALERT: Confe;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1675656890] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1675656898] HOST ALERT: FW Confe;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1675656910] HOST ALERT: FW Cass Entel;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1675656926] HOST ALERT: ADL Diagnostic;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1675656946] HOST ALERT: Confe;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1675656980] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1675656988] HOST ALERT: FW Confe;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1675657000] HOST ALERT: FW Cass Entel;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1675657016] HOST ALERT: ADL Diagnostic;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1675657036] HOST ALERT: Confe;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1675657070] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1675657078] HOST ALERT: FW Confe;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1675657090] HOST ALERT: FW Cass Entel;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1675657106] HOST ALERT: ADL Diagnostic;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1675657126] HOST ALERT: Confe;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1675657160] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1675657168] HOST ALERT: FW Confe;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1675657180] HOST ALERT: FW Cass Entel;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1675657196] HOST ALERT: ADL Diagnostic;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1675657216] HOST ALERT: Confe;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1675659825] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.20, 0.06, 0.03
[1675659945] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.24, 0.14, 0.06
[1675660065] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;3;OK - load average per CPU: 0.15, 0.14, 0.07
[1675664345] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7943.31 MB (84%) - free: 1528.19 MB (16%)
[1675664465] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7643.37 MB (81%) - free: 1828.13 MB (19%)
[1675664585] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7903.24 MB (83%) - free: 1568.27 MB (17%)
[1675664945] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9471.50 MB - used: 7532.66 MB (80%) - free: 1938.84 MB (20%)
[1675665005] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7951.92 MB (84%) - free: 1519.58 MB (16%)
[1675665125] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7792.94 MB (82%) - free: 1678.57 MB (18%)
[1675665245] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7965.39 MB (84%) - free: 1506.11 MB (16%)
[1675665425] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1675665425] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9471.50 MB - used: 7517.94 MB (79%) - free: 1953.56 MB (21%)
[1675665485] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7939.00 MB (84%) - free: 1532.50 MB (16%)
[1675665605] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7680.41 MB (81%) - free: 1791.09 MB (19%)
[1675665725] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7967.81 MB (84%) - free: 1503.69 MB (16%)
[1675665905] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9471.50 MB - used: 7453.48 MB (79%) - free: 2018.02 MB (21%)
[1675665965] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 8013.06 MB (85%) - free: 1458.44 MB (15%)
[1675666085] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7829.00 MB (83%) - free: 1642.50 MB (17%)
[1675666205] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7928.78 MB (84%) - free: 1542.72 MB (16%)
[1675666371] SERVICE ALERT: FW Confe;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 31.35 ms
[1675666372] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 37%, RTA = 29.36 ms
[1675666372] SERVICE ALERT: Confe;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 31.26 ms
[1675666373] SERVICE ALERT: ADL Diagnostic;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 80%, RTA = 28.58 ms
[1675666373] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 58%, RTA = 29.50 ms
[1675666373] HOST ALERT: FW Cass Entel;UP;HARD;1;PING OK - Packet loss = 64%, RTA = 29.54 ms
[1675666373] HOST ALERT: FW Confe;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 31.25 ms
[1675666376] HOST ALERT: FW ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 29.52 ms
[1675666376] HOST ALERT: Confe;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 31.90 ms
[1675666380] HOST ALERT: ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 28.60 ms
[1675666405] SERVICE ALERT: FW Confe;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.28 ms
[1675666411] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 32.96 ms
[1675667045] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9471.50 MB - used: 7560.39 MB (80%) - free: 1911.11 MB (20%)
[1675667225] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7602.42 MB (80%) - free: 1869.08 MB (20%)
[1675667345] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7529.94 MB (80%) - free: 1941.57 MB (20%)
[1675667825] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1675686834] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1675686956] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1675687077] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1675690865] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7607.38 MB (80%) - free: 1864.12 MB (20%)
[1675690985] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7613.00 MB (80%) - free: 1858.51 MB (20%)
[1675691105] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7614.33 MB (80%) - free: 1857.18 MB (20%)
[1675691440] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675691571] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675692305] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9471.50 MB - used: 7516.20 MB (79%) - free: 1955.30 MB (21%)
[1675692485] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7651.12 MB (81%) - free: 1820.38 MB (19%)
[1675692605] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7482.07 MB (79%) - free: 1989.44 MB (21%)
[1675692725] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7659.73 MB (81%) - free: 1811.77 MB (19%)
[1675692845] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7601.26 MB (80%) - free: 1870.24 MB (20%)
[1675692965] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;3;Memory usage: total:9471.50 MB - used: 7273.62 MB (77%) - free: 2197.89 MB (23%)
[1675693085] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7638.88 MB (81%) - free: 1832.62 MB (19%)
[1675693205] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7514.41 MB (79%) - free: 1957.09 MB (21%)
[1675693325] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7695.46 MB (81%) - free: 1776.05 MB (19%)
[1675693445] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7251.62 MB (77%) - free: 2219.89 MB (23%)
[1675693686] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7700.76 MB (81%) - free: 1770.75 MB (19%)
[1675693806] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7226.94 MB (76%) - free: 2244.56 MB (24%)
[1675693986] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7725.71 MB (82%) - free: 1745.79 MB (18%)
[1675694106] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9471.50 MB - used: 7231.28 MB (76%) - free: 2240.23 MB (24%)
[1675694286] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9471.50 MB - used: 7698.27 MB (81%) - free: 1773.23 MB (19%)
[1675694406] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9471.50 MB - used: 7695.57 MB (81%) - free: 1775.93 MB (19%)
[1675694526] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9471.50 MB - used: 7695.56 MB (81%) - free: 1775.94 MB (19%)
[1675695426] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:10434.99 MB - used: 8267.83 MB (79%) - free: 2167.16 MB (21%)
[1675695486] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:10199.35 MB - used: 8258.59 MB (81%) - free: 1940.75 MB (19%)
[1675695606] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:10199.35 MB - used: 8201.49 MB (80%) - free: 1997.86 MB (20%)
[1675695726] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;3;Memory usage: total:10071.34 MB - used: 7951.75 MB (79%) - free: 2119.59 MB (21%)
[1675695846] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:10071.34 MB - used: 8139.21 MB (81%) - free: 1932.13 MB (19%)
[1675695966] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:10071.34 MB - used: 8197.54 MB (81%) - free: 1873.80 MB (19%)
[1675696086] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9815.34 MB - used: 8113.61 MB (83%) - free: 1701.73 MB (17%)
[1675696326] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9815.34 MB - used: 7547.77 MB (77%) - free: 2267.57 MB (23%)
[1675696446] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9687.34 MB - used: 7986.97 MB (82%) - free: 1700.37 MB (18%)
[1675696566] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9687.34 MB - used: 7972.67 MB (82%) - free: 1714.67 MB (18%)
[1675696686] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1675696686] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9687.34 MB - used: 7873.35 MB (81%) - free: 1813.99 MB (19%)
[1675696926] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9687.34 MB - used: 7688.39 MB (79%) - free: 1998.95 MB (21%)
[1675697046] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9687.34 MB - used: 7847.60 MB (81%) - free: 1839.74 MB (19%)
[1675697166] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9687.34 MB - used: 7828.03 MB (81%) - free: 1859.31 MB (19%)
[1675697286] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9687.34 MB - used: 7938.95 MB (82%) - free: 1748.39 MB (18%)
[1675697526] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9687.34 MB - used: 7712.84 MB (80%) - free: 1974.50 MB (20%)
[1675697766] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 8024.32 MB (83%) - free: 1599.02 MB (17%)
[1675697886] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 7999.03 MB (83%) - free: 1624.31 MB (17%)
[1675698006] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 7979.51 MB (83%) - free: 1643.83 MB (17%)
[1675698966] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1675700886] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7696.00 MB (80%) - free: 1927.34 MB (20%)
[1675701006] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 8000.90 MB (83%) - free: 1622.44 MB (17%)
[1675701126] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 7809.04 MB (81%) - free: 1814.30 MB (19%)
[1675701246] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 7997.50 MB (83%) - free: 1625.84 MB (17%)
[1675701906] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7530.10 MB (78%) - free: 2093.24 MB (22%)
[1675701966] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 8034.85 MB (83%) - free: 1588.49 MB (17%)
[1675702086] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 8028.69 MB (83%) - free: 1594.65 MB (17%)
[1675702206] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1675702206] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 7992.85 MB (83%) - free: 1630.49 MB (17%)
[1675703211] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1675703331] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 743.31KB (0.59%), Average OUT: 915.62KB (0.73%)Total RX: 93.98GBytes, Total TX: 186.86GBytes
[1675703346] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1675703946] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7622.52 MB (79%) - free: 2000.82 MB (21%)
[1675704066] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 7699.96 MB (80%) - free: 1923.38 MB (20%)
[1675704186] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 7977.76 MB (83%) - free: 1645.58 MB (17%)
[1675704306] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 8162.88 MB (85%) - free: 1460.46 MB (15%)
[1675704666] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7522.88 MB (78%) - free: 2100.46 MB (22%)
[1675704726] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 7862.92 MB (82%) - free: 1760.43 MB (18%)
[1675704846] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;SOFT;2;Memory usage: total:9623.34 MB - used: 7580.39 MB (79%) - free: 2042.95 MB (21%)
[1675704966] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 7833.54 MB (81%) - free: 1789.80 MB (19%)
[1675705086] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 7996.21 MB (83%) - free: 1627.13 MB (17%)
[1675705206] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1675705206] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 8040.49 MB (84%) - free: 1582.85 MB (16%)
[1675705266] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7567.76 MB (79%) - free: 2055.58 MB (21%)
[1675705386] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;1;Memory usage: total:9623.34 MB - used: 7821.54 MB (81%) - free: 1801.80 MB (19%)
[1675705506] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;SOFT;2;Memory usage: total:9623.34 MB - used: 8023.69 MB (83%) - free: 1599.65 MB (17%)
[1675705626] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;WARNING;HARD;3;Memory usage: total:9623.34 MB - used: 7711.40 MB (80%) - free: 1911.94 MB (20%)
[1675706706] SERVICE FLAPPING ALERT: SRVADMIN.cass.pmc;Memoria;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1675707666] SERVICE ALERT: SRVADMIN.cass.pmc;Memoria;OK;HARD;1;Memory usage: total:9623.34 MB - used: 7136.63 MB (74%) - free: 2486.71 MB (26%)
[1675708466] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1675708586] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 5.38KB (0.00%), Average OUT: 6.41KB (0.01%)Total RX: 14.66GBytes, Total TX: 17.74GBytes
[1675708982] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1675709102] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 726.28KB (0.58%), Average OUT: 732.97KB (0.59%)Total RX: 97.74GBytes, Total TX: 190.90GBytes
[1675709629] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1675709749] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1675710696] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1675710816] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 13.86KB (0.01%), Average OUT: 12.44KB (0.01%)Total RX: 14.70GBytes, Total TX: 17.76GBytes
[1675711169] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1675711290] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1675712652] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Plugin timed out while executing system call
[1675712772] SERVICE ALERT: FW Cass Telsur;Tiempo Activo;OK;SOFT;2;SNMP OK - Timeticks: (101055082) 11 days, 16:42:30.82
[1675713466] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.20, 0.10, 0.08
[1675713586] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.06, 0.09, 0.08
[1675715849] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675715981] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675717873] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1675717993] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 36.49KB (0.03%), Average OUT: 290.93KB (0.23%)Total RX: 98.55GBytes, Total TX: 193.65GBytes
[1675718301] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675718327] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675718432] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675718465] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1675718620] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1675718740] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1675719465] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1675719586] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1675732154] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1675732274] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1675732394] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1675732694] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 84% (5 min average)
[1675732754] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 67% (5 min average)