[1696058402] SERVICE ALERT: SRV_SOS2021.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.0.16. Check system logs on 10.0.0.16
[1696058522] SERVICE ALERT: SRV_SOS2021.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 191 day(s) 14 hour(s) 59 minute(s)
[1696064213] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;1;YUM CRITICAL: Yum nagios plugin has self terminated after exceeding the timeout (30 seconds)
[1696064303] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;2;YUM CRITICAL: Yum nagios plugin has self terminated after exceeding the timeout (30 seconds)
[1696064365] SERVICE ALERT: SRV_nagios;Actualizacion;OK;SOFT;3;YUM OK: 0 Updates Available
[1696064484] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only -1 bytes received (4 expected).
[1696064604] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4964 MB - 64% [SWAP] Total: 7944 MB - Used: 780 MB - 9%
[1696079369] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 95% (5 min average)
[1696079489] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1696079609] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1696080149] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 87% (5 min average)
[1696080238] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (85.31%) too high
[1696080359] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.26%) too high
[1696080479] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (93.03%) too high
[1696081284] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1696081404] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4960 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696084884] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1696085004] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4960 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696087199] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 583.06KB (0.47%), Average OUT: 105.87MB (84.69%)Total RX: 44404.84GBytes, Total TX: 102887.76GBytes
[1696087259] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (92.10%) too high
[1696087380] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (86.60%) too high
[1696087500] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (87.84%) too high
[1696087740] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 539.59KB (0.43%), Average OUT: 103.56MB (82.85%)Total RX: 44405.16GBytes, Total TX: 102947.34GBytes
[1696087801] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (91.04%) too high
[1696087921] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.85%) too high
[1696088041] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1696088041] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (88.80%) too high
[1696088484] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1696088521] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 568.76KB (0.46%), Average OUT: 106.23MB (84.98%)Total RX: 44405.64GBytes, Total TX: 103035.13GBytes
[1696088581] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (92.64%) too high
[1696088604] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4961 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696088702] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.28%) too high
[1696088822] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.11%) too high
[1696089962] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1696090562] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 7.39MB (5.91%), Average OUT: 101.27MB (81.02%)Total RX: 44407.30GBytes, Total TX: 103266.50GBytes
[1696090623] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (89.38%) too high
[1696090743] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (89.25%) too high
[1696090863] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (91.97%) too high
[1696091223] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 553.45KB (0.44%), Average OUT: 102.95MB (82.36%)Total RX: 44407.69GBytes, Total TX: 103340.64GBytes
[1696091283] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;1;WARNING - OUT bandwidth (93.44%) too high
[1696091404] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;SOFT;2;WARNING - OUT bandwidth (90.98%) too high
[1696091524] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1696091524] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;WARNING;HARD;3;WARNING - OUT bandwidth (92.45%) too high
[1696092665] SERVICE FLAPPING ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1696092724] SERVICE ALERT: Servidores CASS;Ancho de Banda GE/1/0/3;OK;HARD;1;OK - Average IN: 995.44KB (0.80%), Average OUT: 97.34MB (77.87%)Total RX: 44408.74GBytes, Total TX: 103510.56GBytes
[1696095602] SERVICE ALERT: SRV_SOS2021.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.16 port 5666: Connection refused
[1696095722] SERVICE ALERT: SRV_SOS2021.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 192 day(s) 1 hour(s) 19 minute(s)
[1696098084] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1696098204] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4961 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696101684] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1696101804] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4962 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696105284] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1696105404] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4964 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696110684] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1696110804] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4959 MB - 64% [SWAP] Total: 7944 MB - Used: 790 MB - 9%
[1696120284] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.13 port 5666: Connection refused
[1696120404] SERVICE ALERT: SRV11VIBE_cass.cl;Memoria;OK;SOFT;2;[MEMORY] Total: 7693 MB - Used: 4967 MB - 64% [SWAP] Total: 7944 MB - Used: 791 MB - 9%