[1685168401] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1685168521] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1818356 MiB (92.85% inode=99%):
[1685170202] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1685170322] SERVICE ALERT: SRV_email.cass.cl;Procesos Zombies;OK;SOFT;2;PROCS OK: 0 processes with STATE = Z
[1685176801] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25: Connection reset by peer
[1685176921] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1818352 MiB (92.85% inode=99%):
[1685182603] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685182604] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1685182735] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685182736] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685183401] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685183402] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1685183533] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685183534] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1685194142] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1685194143] HOST ALERT: FW Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1685194165] SERVICE ALERT: Electric Works;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1685194191] HOST ALERT: Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1685194221] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1685194255] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1685194292] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1685194311] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1685194315] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1685194345] HOST ALERT: Electric Works;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1685194401] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1685194435] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1685194491] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;5;(Host check timed out after 30.00 seconds)
[1685194525] HOST ALERT: Electric Works;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1685194581] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1685194615] HOST ALERT: Electric Works;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1685194671] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1685194705] HOST ALERT: Electric Works;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1685194761] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1685194795] HOST ALERT: Electric Works;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1685194851] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1685194885] HOST ALERT: Electric Works;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1685194941] HOST ALERT: FW Electric Works;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1685194975] HOST ALERT: Electric Works;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1685195401] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not connect to 10.0.2.25. Check system logs on 10.0.2.25
[1685195522] SERVICE ALERT: SRV_email.cass.cl;Almacenamiento Mail;OK;SOFT;2;DISK OK - free space: /mail 1818340 MiB (92.85% inode=99%):
[1685196507] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;1;CPU Load 96% (5 min average)
[1685196627] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;SOFT;2;CPU Load 99% (5 min average)
[1685196748] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1685197588] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1685197708] SERVICE ALERT: Veeam2023.cass.pmc;CPU;WARNING;HARD;3;CPU Load 99% (5 min average)
[1685197768] SERVICE FLAPPING ALERT: Veeam2023.cass.pmc;CPU;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1685197768] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;HARD;1;CPU Load 91% (5 min average)
[1685198908] SERVICE FLAPPING ALERT: Veeam2023.cass.pmc;CPU;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1685200407] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;OK;HARD;1;OK - no pending updates.
[1685223569] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 33.61 ms
[1685223573] HOST ALERT: Electric Works;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 33.63 ms
[1685223578] HOST ALERT: FW Electric Works;UP;HARD;1;PING OK - Packet loss = 70%, RTA = 32.38 ms
[1685223606] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 32.15 ms
[1685223651] HOST ALERT: FW Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1685223655] SERVICE ALERT: Electric Works;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1685223659] HOST ALERT: Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1685223689] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1685223692] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1685223745] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1685223779] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1685223805] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1685223835] HOST ALERT: Electric Works;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1685223869] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1685223925] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1685223959] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1685224015] HOST ALERT: Electric Works;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1685224049] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1685224105] HOST ALERT: Electric Works;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1685224139] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1685224195] HOST ALERT: Electric Works;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1685224229] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1685224285] HOST ALERT: Electric Works;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1685224319] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1685224375] HOST ALERT: Electric Works;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1685224409] HOST ALERT: FW Electric Works;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1685224465] HOST ALERT: Electric Works;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1685231189] HOST ALERT: Electric Works;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 31.70 ms
[1685231216] HOST ALERT: FW Electric Works;UP;HARD;1;PING WARNING - Packet loss = 81%, RTA = 33.11 ms
[1685231220] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 31.95 ms
[1685231226] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 32.83 ms
[1685232602] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1685232722] SERVICE ALERT: SRV_email.cass.cl;Correo Smtp;OK;SOFT;2;SMTP OK - 0.025 sec. response time
[1685241001] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTP;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.0.92 port 5666: Connection refused
[1685241122] SERVICE ALERT: SRV_chat.cass.cl;Puerto HTTP;OK;SOFT;2;TCP OK - 0.000 second response time on port 80
[1685243679] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1685243800] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1685243921] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional