[1678510588] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 368 processes
[1678510708] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 323 processes
[1678511098] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.35, 0.12, 0.05
[1678511218] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;2;CRITICAL - load average per CPU: 0.43, 0.23, 0.10
[1678511338] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.19, 0.21, 0.11
[1678511578] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.02, 0.12, 0.10
[1678515478] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.14, 0.14, 0.10
[1678515598] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1678515718] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1678524280] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678524287] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678524288] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678524291] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678524412] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678524418] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678524420] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678524422] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678531439] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.44, 0.24, 0.18
[1678531919] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.01, 0.17, 0.19
[1678532579] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.01, 0.03, 0.10
[1678537402] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537408] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537412] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537534] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537539] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537544] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537684] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537689] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537694] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537711] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1678537816] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537820] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537826] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678537843] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1678560268] SERVICE ALERT: Confe;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1678560268] SERVICE ALERT: FW Confe;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1678560273] HOST ALERT: Confe;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678560288] HOST ALERT: FW Confe;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1678560318] HOST ALERT: Confe;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1678560363] HOST ALERT: FW Confe;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1678560408] HOST ALERT: Confe;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1678560418] SERVICE ALERT: Confe;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1678560418] SERVICE ALERT: FW Confe;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1678560453] HOST ALERT: FW Confe;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1678560498] HOST ALERT: Confe;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1678560543] HOST ALERT: FW Confe;UNREACHABLE;SOFT;4;(Host check timed out after 30.01 seconds)
[1678560588] HOST ALERT: Confe;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1678560633] HOST ALERT: FW Confe;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1678560678] HOST ALERT: Confe;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1678560723] HOST ALERT: FW Confe;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1678560768] HOST ALERT: Confe;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1678560814] HOST ALERT: FW Confe;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1678560858] HOST ALERT: Confe;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1678560904] HOST ALERT: FW Confe;UNREACHABLE;SOFT;8;(Host check timed out after 30.01 seconds)
[1678560948] HOST ALERT: Confe;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1678560994] HOST ALERT: FW Confe;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1678561038] HOST ALERT: Confe;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1678561084] HOST ALERT: FW Confe;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1678563008] HOST ALERT: FW Confe;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 37.08 ms
[1678563022] HOST ALERT: Confe;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 37.28 ms
[1678563032] SERVICE ALERT: Confe;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 37.69 ms
[1678563033] SERVICE ALERT: FW Confe;Ping;OK;HARD;1;PING OK - Packet loss = 16%, RTA = 37.17 ms
[1678567220] HOST ALERT: FW ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678567220] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678567220] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1678567223] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1678567248] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1678567248] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678567250] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1678567310] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1678567338] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1678567340] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1678567398] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1678567400] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1678567428] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1678567430] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;4;(Host check timed out after 30.00 seconds)
[1678567490] HOST ALERT: ADL Diagnostic;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1678567518] HOST ALERT: FW Cass Entel;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1678567520] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;5;(Host check timed out after 30.01 seconds)
[1678567580] HOST ALERT: ADL Diagnostic;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1678567608] HOST ALERT: FW Cass Entel;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1678567610] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;6;(Host check timed out after 30.01 seconds)
[1678567670] HOST ALERT: ADL Diagnostic;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1678567698] HOST ALERT: FW Cass Entel;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1678567700] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;7;(Host check timed out after 30.01 seconds)
[1678567760] HOST ALERT: ADL Diagnostic;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1678567788] HOST ALERT: FW Cass Entel;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1678567790] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;8;(Host check timed out after 30.00 seconds)
[1678567850] HOST ALERT: ADL Diagnostic;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1678567878] HOST ALERT: FW Cass Entel;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1678567880] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;9;(Host check timed out after 30.01 seconds)
[1678567940] HOST ALERT: ADL Diagnostic;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1678567968] HOST ALERT: FW Cass Entel;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1678567970] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;HARD;10;(Host check timed out after 30.01 seconds)
[1678568030] HOST ALERT: ADL Diagnostic;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1678568058] HOST ALERT: FW Cass Entel;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1678569605] HOST ALERT: FW ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 68%, RTA = 35.59 ms
[1678569606] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 70%, RTA = 35.05 ms
[1678569606] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 64%, RTA = 34.11 ms
[1678569609] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 76%, RTA = 33.64 ms
[1678569609] HOST ALERT: ADL Diagnostic;UP;HARD;1;PING OK - Packet loss = 75%, RTA = 34.54 ms
[1678569613] HOST ALERT: FW Cass Entel;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 33.61 ms
[1678570460] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1678570460] HOST ALERT: FW ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678570460] HOST ALERT: ADL Diagnostic;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678570460] SERVICE ALERT: ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1678570462] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 1.56 ms
[1678570463] SERVICE ALERT: FW ADL Diagnostic;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1678570479] HOST ALERT: FW Cass Entel;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1678570490] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1678570550] HOST ALERT: FW Cass Entel;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1678570550] HOST ALERT: ADL Diagnostic;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1678570580] HOST ALERT: FW ADL Diagnostic;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1678570586] SERVICE ALERT: SRV_Oraculo_cass.cl;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.07 ms
[1678570610] SERVICE ALERT: FW Cass Entel;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1678570640] HOST ALERT: FW Cass Entel;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1678570640] HOST ALERT: ADL Diagnostic;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1678570664] HOST ALERT: FW ADL Diagnostic;UP;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 42.97 ms
[1678570666] SERVICE ALERT: FW Cass Entel;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 81%, RTA = 43.34 ms
[1678570675] SERVICE ALERT: ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 38.09 ms
[1678570677] SERVICE ALERT: FW ADL Diagnostic;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 36.36 ms
[1678570679] HOST ALERT: ADL Diagnostic;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 36.82 ms
[1678570704] HOST ALERT: FW Cass Entel;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 34.87 ms
[1678570704] SERVICE FLAPPING ALERT: FW Cass Entel;Ping;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1678570704] SERVICE ALERT: FW Cass Entel;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 35.43 ms
[1678571844] SERVICE FLAPPING ALERT: FW Cass Entel;Ping;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)