[1714277488] HOST ALERT: FW Pasmar;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1714277491] SERVICE ALERT: FW Pasmar;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1714277493] SERVICE ALERT: Pasmar;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 1.20 ms
[1714277578] HOST ALERT: FW Pasmar;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1714277617] SERVICE ALERT: Pasmar;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1714277642] HOST ALERT: FW Pasmar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1714277645] SERVICE ALERT: FW Pasmar;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 1.53 ms
[1714284604] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 10.0.2.25 port 5666: Connection refused
[1714284724] SERVICE ALERT: SRV_email.cass.cl;Tiempo Activo;OK;SOFT;2;Uptime OK: 50 day(s) 18 hour(s) 50 minute(s)
[1714287165] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1714287165] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1714287166] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1714287242] SERVICE ALERT: SRV_GLPI;CPU;WARNING;SOFT;1;WARNING: Percent was 65.00 %
[1714287297] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287298] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287298] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287363] SERVICE ALERT: SRV_GLPI;CPU;OK;SOFT;2;OK: Percent was 0.00 %
[1714287601] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1714287601] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1714287602] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1714287630] HOST ALERT: FW Gore Los Lagos;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1714287633] SERVICE ALERT: FW Gore Los Lagos;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1714287674] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;HARD;1;PING OK - Packet loss = 58%, RTA = 1.26 ms
[1714287675] SERVICE ALERT: Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 80%, RTA = 1.18 ms
[1714287678] HOST ALERT: FW Gore Los Lagos;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1714287733] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287734] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287734] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1714287799] SERVICE ALERT: Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.23 ms
[1714306041] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1714306042] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1714306047] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1714306051] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1714306051] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1714306067] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;HARD;3;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306067] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306067] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1714306128] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 102.01 Gb (85%) - free 17.39 Gb (15%)
[1714306128] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;HARD;1;Memory usage: total:23551.21 MB - used: 1951.56 MB (8%) - free: 21599.66 MB (92%)
[1714306141] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1714306161] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 227.51 Gb (22%) - free 796.47 Gb (78%)
[1714306162] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;SOFT;2;CPU Load 0% (5 min average)
[1714306167] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1714306197] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1714306198] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1714306198] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1714306198] SERVICE ALERT: SRV01USER_cass.cl;Ping;CRITICAL;HARD;1;CRITICAL - Host Unreachable (10.0.0.12)
[1714306204] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.12 port 5666: No route to host
[1714306243] SERVICE ALERT: SRV01USER_cass.cl;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 2.44 ms
[1714306247] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.85 ms
[1714306259] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Almacenamiento C;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1714306259] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 101.56 Gb (85%) - free 17.84 Gb (15%)
[1714306287] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306289] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306289] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;SOFT;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306317] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;SOFT;2;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306319] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;HARD;3;connect to address 10.0.0.12 and port 12489: Connection refused
[1714306336] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1714306336] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.
[1714306380] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;WARNING;HARD;3;c:\ - total: 119.40 Gb - used: 101.56 Gb (85%) - free 17.84 Gb (15%)
[1714306407] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1714306409] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;SOFT;2;u:\ - total: 1023.98 Gb - used: 227.51 Gb (22%) - free 796.47 Gb (78%)
[1714306409] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;SOFT;2;CPU Load 1% (5 min average)
[1714306437] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;SOFT;3;Memory usage: total:23551.21 MB - used: 1991.88 MB (8%) - free: 21559.34 MB (92%)
[1714307404] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1714307520] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Almacenamiento C;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1714327083] SERVICE ALERT: SRV_GLPI;CPU;CRITICAL;SOFT;1;CRITICAL: Percent was 99.50 %
[1714327204] SERVICE ALERT: SRV_GLPI;CPU;OK;SOFT;2;OK: Percent was 0.00 %
[1714354252] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1714354374] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1714363084] SERVICE ALERT: SRV_GLPI;CPU;CRITICAL;SOFT;1;CRITICAL: Percent was 89.50 %