[1676437476] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.35, 0.12, 0.05
[1676437596] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.27, 0.18, 0.08
[1676437717] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.15, 0.17, 0.09
[1676437837] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.06, 0.13, 0.09
[1676459168] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;1;YUM CRITICAL: 2 Updates Available
[1676459231] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;2;YUM CRITICAL: 2 Updates Available
[1676459294] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;SOFT;3;YUM CRITICAL: 2 Updates Available
[1676459357] SERVICE ALERT: SRV_nagios;Actualizacion;CRITICAL;HARD;4;YUM CRITICAL: 2 Updates Available
[1676464819] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 3 critical, 0 optional
[1676464940] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 3 critical, 0 optional
[1676465062] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 3 critical, 0 optional
[1676467507] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676467576] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676467642] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;(No output on stdout) stderr:
[1676467702] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676467706] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1676467822] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 7.10KB (0.01%), Average OUT: 7.80KB (0.01%)Total RX: 25.01GBytes, Total TX: 23.95GBytes
[1676467827] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 5.68MB (4.55%), Average OUT: 3.17MB (2.54%)Total RX: 8071.32GBytes, Total TX: 3790.76GBytes
[1676468377] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1676477853] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676477973] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 6.76KB (0.01%), Average OUT: 7.58KB (0.01%)Total RX: 25.49GBytes, Total TX: 24.17GBytes
[1676478457] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1676478577] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 88.32B (0.00%), Average OUT: 16.42B (0.00%)Total RX: 8087.54GBytes, Total TX: 3799.16GBytes
[1676480673] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1676480793] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1676662419] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;SOFT;1;DISK UNKNOWN - free space:
[1676662449] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1676662510] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;HARD;1;OK - no pending updates.
[1676662539] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;SOFT;2;DISK UNKNOWN - free space:
[1676662659] SERVICE ALERT: SRV_chat.cass.cl;Almacenamiento /;UNKNOWN;HARD;3;DISK UNKNOWN - free space:
[1676667008] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1676667128] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1676667782] SERVICE ALERT: Clinica PV Movistar;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1676667784] SERVICE ALERT: FW Clinica PV Movistar;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1676667794] HOST ALERT: Clinica PV Movistar;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1676667797] HOST ALERT: Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1676667806] HOST ALERT: FW Clinica PV Movistar;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1676667807] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 96%, RTA = 7406.63 ms
[1676667809] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1676667810] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1676667836] HOST ALERT: Clinica PV Movistar;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1676667840] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1676667884] HOST ALERT: FW Clinica PV Movistar;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1676667899] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1676667926] HOST ALERT: Clinica PV Movistar;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1676667930] HOST ALERT: Electric Works;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1676667932] SERVICE ALERT: Clinica PV Movistar;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676667934] SERVICE ALERT: FW Clinica PV Movistar;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676667959] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676667974] HOST ALERT: FW Clinica PV Movistar;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1676667989] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.01 seconds)
[1676668016] HOST ALERT: Clinica PV Movistar;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1676668019] SERVICE ALERT: FW Electric Works;Ping;WARNING;HARD;3;PING WARNING - Packet loss = 93%, RTA = 101.80 ms
[1676668020] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1676668021] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 94.11 ms
[1676668025] HOST ALERT: Electric Works;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 94.01 ms
[1676668029] HOST ALERT: FW Electric Works;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 101.01 ms
[1676668042] SERVICE ALERT: FW Clinica PV Movistar;Ping;OK;HARD;1;PING OK - Packet loss = 73%, RTA = 100.96 ms
[1676668042] SERVICE ALERT: Clinica PV Movistar;Ping;OK;HARD;1;PING OK - Packet loss = 76%, RTA = 96.90 ms
[1676668043] HOST ALERT: FW Clinica PV Movistar;UP;SOFT;1;PING OK - Packet loss = 50%, RTA = 100.98 ms
[1676668047] HOST ALERT: Clinica PV Movistar;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 97.06 ms
[1676668053] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 101.97 ms
[1676675126] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1676675166] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1676675246] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1676676508] SERVICE ALERT: Electric Works;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 96%, RTA = 99.52 ms
[1676676511] HOST ALERT: Electric Works;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1676676515] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;1;(Host check timed out after 30.01 seconds)
[1676676539] SERVICE ALERT: FW Electric Works;Ping;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1676676545] HOST ALERT: Electric Works;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1676676601] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;2;(Host check timed out after 30.01 seconds)
[1676676635] HOST ALERT: Electric Works;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1676676658] SERVICE ALERT: Electric Works;Ping;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1676676691] HOST ALERT: FW Electric Works;UNREACHABLE;SOFT;3;(Host check timed out after 30.00 seconds)
[1676676725] HOST ALERT: Electric Works;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1676676752] SERVICE ALERT: Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 99.79 ms
[1676676755] HOST ALERT: FW Electric Works;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 102.68 ms
[1676676755] SERVICE ALERT: FW Electric Works;Ping;OK;HARD;1;PING OK - Packet loss = 28%, RTA = 102.89 ms
[1676676756] HOST ALERT: Electric Works;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 99.92 ms
[1676676981] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1676677716] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;UNKNOWN;SOFT;1;NSClient - ERROR: No performance data from command: check_drivesize
[1676677811] SERVICE ALERT: SRV01USER_cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1676677814] HOST ALERT: SRV01USER_cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.12)
[1676677819] SERVICE ALERT: SRV01USER_cass.cl;Memoria;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1676677821] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.12 port 5666: Connection refused
[1676677836] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;CRITICAL;HARD;3;connect to address 10.0.0.12 and port 12489: Connection refused
[1676677839] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1676677839] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;CRITICAL;HARD;1;connect to address 10.0.0.12 and port 12489: Connection refused
[1676677844] HOST ALERT: SRV01USER_cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1.88 ms
[1676677880] SERVICE ALERT: SRV01USER_cass.cl;Memoria;OK;HARD;1;Memory usage: total:23551.21 MB - used: 1577.56 MB (7%) - free: 21973.65 MB (93%)
[1676677896] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento C;OK;HARD;1;c:\ - total: 119.40 Gb - used: 82.03 Gb (69%) - free 37.37 Gb (31%)
[1676677899] SERVICE ALERT: SRV01USER_cass.cl;Almacenamiento U;OK;HARD;1;u:\ - total: 1023.98 Gb - used: 481.72 Gb (47%) - free 542.26 Gb (53%)
[1676677899] SERVICE ALERT: SRV01USER_cass.cl;Tiempo Activo;OK;HARD;1;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1676677911] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1676677931] SERVICE ALERT: SRV01USER_cass.cl;CPU;OK;SOFT;2;CPU Load 0% (5 min average)
[1676678095] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1676678095] SERVICE ALERT: SRV01USER_cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1676678962] SERVICE ALERT: SRVCONT.cass.cl;Almacenamiento;UNKNOWN;SOFT;1;NSClient - ERROR: No performance data from command: check_drivesize
[1676678981] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;WARNING;HARD;3;updates installed, reboot required
[1676678989] SERVICE ALERT: SRVCONT.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1676678989] SERVICE ALERT: SRVCONT.cass.cl;Memoria;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1676678998] HOST ALERT: SRVCONT.cass.cl;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.95)
[1676679006] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 10.0.0.95 port 5666: Connection refused
[1676679012] HOST ALERT: SRVCONT.cass.cl;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 2.36 ms
[1676679022] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;CRITICAL;SOFT;1;connect to address 10.0.0.95 and port 12489: Connection refused
[1676679083] SERVICE ALERT: SRVCONT.cass.cl;Almacenamiento;OK;SOFT;2;c:\ - total: 299.33 Gb - used: 38.52 Gb (13%) - free 260.81 Gb (87%)
[1676679109] SERVICE ALERT: SRVCONT.cass.cl;Memoria;OK;SOFT;2;Memory usage: total:14143.13 MB - used: 1773.37 MB (13%) - free: 12369.76 MB (87%)
[1676679109] SERVICE ALERT: SRVCONT.cass.cl;CPU;OK;SOFT;2;CPU Load 7% (5 min average)
[1676679117] SERVICE FLAPPING ALERT: SRV01USER_cass.cl;Actualizacion;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1676679142] SERVICE ALERT: SRVCONT.cass.cl;Tiempo Activo;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1676679216] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;UNKNOWN;HARD;3;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1676679279] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1676688699] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;UNKNOWN;SOFT;1;NSClient - ERROR: No performance data from command: check_drivesize
[1676688743] SERVICE ALERT: Veeam2023.cass.pmc;CPU;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1676688791] SERVICE ALERT: Veeam2023.cass.pmc;Actualizacion;CRITICAL;HARD;3;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676688819] SERVICE ALERT: Veeam2023.cass.pmc;Almacenamiento;OK;SOFT;2;c:\ - total: 99.33 Gb - used: 49.00 Gb (49%) - free 50.33 Gb (51%)
[1676688863] SERVICE ALERT: Veeam2023.cass.pmc;CPU;OK;SOFT;2;CPU Load 7% (5 min average)