[1677814406] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 388 processes
[1677814526] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 329 processes
[1677819898] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.38, 0.13, 0.05
[1677820018] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;2;CRITICAL - load average per CPU: 0.33, 0.22, 0.10
[1677820138] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.16, 0.19, 0.10
[1677820318] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.04, 0.13, 0.10
[1677844559] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677844612] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677844684] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 7.20MB (5.76%), Average OUT: 3.95MB (3.16%)Total RX: 14446.64GBytes, Total TX: 6392.93GBytes
[1677844732] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 35.09KB (0.03%), Average OUT: 408.71KB (0.33%)Total RX: 622.43GBytes, Total TX: 551.47GBytes
[1677853010] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677853095] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677853130] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1677853225] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;OK - Average IN: 2.05MB (1.64%), Average OUT: 1.64MB (1.31%)Total RX: 14500.80GBytes, Total TX: 6421.26GBytes
[1677855902] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;UNKNOWN;SOFT;1;UNKNOWN: Improper interface status, status received: DNE - No response from remote host "201.220.96.34"
[1677856022] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 44.92KB (0.04%), Average OUT: 937.25KB (0.75%)Total RX: 635.50GBytes, Total TX: 555.61GBytes
[1677862802] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677862809] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677862848] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677862848] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1677862934] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677862940] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677862979] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677862981] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1677863035] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;HARD;3;The command (check_updates) returned an invalid return code: 255
[1677863096] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1677863374] SERVICE ALERT: FW Gore Los Lagos;Ping;WARNING;SOFT;1;PING WARNING - Packet loss = 83%, RTA = 7.12 ms
[1677863498] SERVICE ALERT: FW Gore Los Lagos;Ping;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 9.59 ms
[1677863895] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677864020] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1677864725] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677864850] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1677865999] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677866119] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 8.80KB (0.01%), Average OUT: 12.05KB (0.01%)Total RX: 44.43GBytes, Total TX: 34.82GBytes
[1677866850] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677866892] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677866970] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1677867017] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;OK - Average IN: 45.68KB (0.04%), Average OUT: 375.82KB (0.30%)Total RX: 636.36GBytes, Total TX: 560.50GBytes
[1677867567] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677867569] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677867576] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677867700] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677867702] SERVICE ALERT: FW Cass Telsur;Ancho de Banda LAN;OK;SOFT;2;(No output on stdout) stderr:
[1677867705] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;(No output on stdout) stderr:
[1677867711] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;2;(No output on stdout) stderr:
[1677867800] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677867820] SERVICE ALERT: FW Cass Telsur;Estado Puerto DMZ;OK;SOFT;2;OK: Interface eth2 (index 4) is up.
[1677867895] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677867900] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677867920] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1677868021] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677868025] SERVICE ALERT: FW Cass Telsur;Estado Puerto WAN;OK;SOFT;2;OK: Interface eth1 (index 3) is up.
[1677868030] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;(No output on stdout) stderr:
[1677868151] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;CRITICAL;SOFT;2;(No output on stdout) stderr:
[1677868160] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677868277] SERVICE ALERT: FW Cass Telsur;Ancho de Banda WAN;OK;SOFT;3;OK - Average IN: 14.54B (0.00%), Average OUT: 23.28B (0.00%)Total RX: 14502.30GBytes, Total TX: 6423.82GBytes
[1677868280] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 21.93KB (0.02%), Average OUT: 7.28KB (0.01%)Total RX: 44.45GBytes, Total TX: 34.83GBytes
[1677870570] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;WARNING;SOFT;1;WARNING: SNMP error: No response from remote host "201.220.96.34"
[1677870690] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;CRITICAL;SOFT;1;(No output on stdout) stderr:
[1677870690] SERVICE ALERT: FW Cass Telsur;Estado Puerto LAN;OK;SOFT;2;OK: Interface eth0 (index 2) is up.
[1677870810] SERVICE ALERT: FW Cass Telsur;Ancho de Banda DMZ;OK;SOFT;2;OK - Average IN: 5.15KB (0.00%), Average OUT: 7.01KB (0.01%)Total RX: 44.51GBytes, Total TX: 34.85GBytes
[1677873441] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;HARD;1;OK - no pending updates.
[1677875605] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 356 processes
[1677875725] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 356 processes
[1677875845] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 359 processes
[1677875906] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 334 processes
[1677875965] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 357 processes
[1677876085] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 353 processes
[1677876205] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 354 processes
[1677876265] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Procesos;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1677876265] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 335 processes
[1677876325] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 354 processes
[1677876445] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 361 processes
[1677876565] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 358 processes
[1677876626] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 337 processes
[1677876685] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 356 processes
[1677876805] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 354 processes
[1677876925] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;3;PROCS OK: 350 processes
[1677877165] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 353 processes
[1677877286] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 355 processes
[1677877406] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 354 processes
[1677877466] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 338 processes
[1677877526] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1677877646] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 352 processes
[1677877766] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 351 processes
[1677877826] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 331 processes
[1677877886] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 352 processes
[1677878007] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 351 processes
[1677878127] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 359 processes
[1677878187] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 338 processes
[1677878247] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 355 processes
[1677878367] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;2;PROCS WARNING: 354 processes
[1677878487] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;HARD;3;PROCS WARNING: 353 processes
[1677878547] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;HARD;1;PROCS OK: 338 processes
[1677878727] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 352 processes
[1677878847] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 349 processes
[1677879807] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;WARNING;SOFT;1;PROCS WARNING: 356 processes
[1677879927] SERVICE ALERT: SRV_SOS2021.cass.cl;Procesos;OK;SOFT;2;PROCS OK: 349 processes
[1677880047] SERVICE FLAPPING ALERT: SRV_SOS2021.cass.cl;Procesos;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1677889643] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;OK;HARD;1;OK - no pending updates.