[1676091888] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.34, 0.12, 0.05
[1676092008] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.18, 0.14, 0.07
[1676095548] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.21, 0.10, 0.05
[1676095668] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;SOFT;2;OK - load average per CPU: 0.14, 0.12, 0.06
[1676096328] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;1;WARNING - load average per CPU: 0.14, 0.14, 0.10
[1676096448] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;SOFT;2;WARNING - load average per CPU: 0.15, 0.14, 0.11
[1676096568] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.14, 0.14, 0.11
[1676112199] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676112216] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676112221] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676112331] SERVICE ALERT: SRV_SOS2021.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1676112347] SERVICE ALERT: SRV_chat.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1676112353] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1676115708] SERVICE ALERT: SRV_email.cass.cl;CPU;CRITICAL;HARD;3;CRITICAL - load average per CPU: 0.32, 0.18, 0.15
[1676116188] SERVICE ALERT: SRV_email.cass.cl;CPU;WARNING;HARD;3;WARNING - load average per CPU: 0.02, 0.18, 0.19
[1676116848] SERVICE ALERT: SRV_email.cass.cl;CPU;OK;HARD;1;OK - load average per CPU: 0.01, 0.03, 0.10
[1676119070] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;1;Updates: 1 critical, 0 optional
[1676119191] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;SOFT;2;Updates: 1 critical, 0 optional
[1676119312] SERVICE ALERT: SRVCONT.cass.cl;Actualizacion;CRITICAL;HARD;3;Updates: 1 critical, 0 optional
[1676151907] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;UNKNOWN;SOFT;1;CHECK_NRPE: Receive header underflow - only 0 bytes received (4 expected).
[1676152028] SERVICE ALERT: SRVADMIN.cass.pmc;Actualizacion;OK;SOFT;2;OK - no pending updates.
[1676168963] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676168963] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1676169095] SERVICE ALERT: SRV_email.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available
[1676169096] SERVICE ALERT: SRV_Cfiler.cass.cl;Actualizacion;OK;SOFT;2;Updates OK : no updates available